```python
# Initialize Otter
import otter
grader = otter.Notebook("hw10.ipynb")
```

## Name: [Your name] and (your onyen)

I encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 10: kNN and cross validation (32 points)
  
**Submission instruction**: Please create a zip file and a pdf via File -> Print (or cmd + P on mac), and upload it to Gradescope.   Even if you decide to finish it at home, please submit what you have by the end of class to make sure you get some credit!


## Problem 1a (10 points)
A team of marine biologists from Antarctica recently discovered group of penguins floating in the sea on an ice floe.  They have no idea where they are from or what species they are, so the scientists are depending on you for help.  Your mission is to help them decide what species the penguins are from **given only measurements of their culmen length and depth**.  Your approach to the problem will be to use the Palmer penguin dataset to build the best darn kNN classifier that you can using 5-fold cross validation to decide its hyperparameter.  Use as many code cells as you need, and use a plot to show results that let you choose the hyperparameter.  Good luck!

```python
import pandas as pd
import seaborn as sns
from matplotlib import pyplot as plt
from sklearn.neighbors import KNeighborsClassifier
penguins = ...
...


```

```python
grader.check("penguins")
```

## Problem 1b (1 point)
What did you choose for your hyperparameter?

*Your answer here*

## Problem 1c (2 points)
The table below are the measurements of the penguins found on the ice floe.  Using your decided upon hyperparameter, train a kNN classifier using culmen length and culmen depth from the entire dataset to predict the species of each ice floe penguin using sklearn.

```python
penguins_ice_flow=pd.DataFrame([[50.1, 19.7],[49.1, 17.0],[48.6, 18.9]], columns=X.columns)
# predict the species of the penguins on the ice flow using your best kNN model
...

```

## Problem 2a (10 points)
It turns out the penguins from the ice floe love playing on your portable scale!  So now you also have body mass measurements!  Pick a hyperparemeter using 5-fold cross validation for a kNN model that uses **culmen length and culmen depth and body mass** to predict species.  Again visually show how you select the hyperparameter. 

```python
import pandas as pd
import seaborn as sns
from matplotlib import pyplot as plt
from sklearn.neighbors import KNeighborsClassifier
# write your code here. Use as many code cells as you need.
...


```

## Problem 2b (1 point)
What did you decide for your hyperparameter?

*Your answer here*

## Problem 2c (2 points)
The table below are the measurements of the penguins found on the ice floe including body mass.  Using your decided upon hyperparameter from 2b, train a kNN classifier using culmen length, culmen depth, and body mass from the entire dataset to predict the species of each ice floe penguin using sklearn.
Use the extra measurements of body mass below to make new predictions for your penguins.
 

```python
penguins_ice_flow=pd.DataFrame([[50.1, 19.7,4900.0],[49.1, 17.0,4600.0],[48.6, 18.9,5000.0]], columns=X.columns)
# predict the species of the penguins on the ice flow using your best kNN model
...
```

## Problem 3a (3 points)
Did you get a consensus?  So it seems like you are done?  Well, if you got a consensus of Gentoo, you might want to remember that kNN uses distances.  And when you use distance metrics, you need to make sure of something about the data.  Adjust for this, and repeat what you did in 2a (do a hyperparameter search using 5-fold cross-validation for a kNN that uses **culmen length and culmen depth and body mass** to predict species).

```python
import pandas as pd
import seaborn as sns
from matplotlib import pyplot as plt
from sklearn.neighbors import KNeighborsClassifier
# write your code here. Use as many code cells as you need.
...


```

## Problem 3b (1 point)
What value did you pick for your hyperparementer?

*Your answer here*

## Problem 3c (2 points)
Finally!  Using your decided upon hyperparameter from 3b, train a kNN classifier using culmen length, culmen depth, and body mass from the entire dataset to predict the species of each ice floe penguin using sklearn.

```python
penguins_ice_flow=pd.DataFrame([[50.1, 19.7,4900.0],[49.1, 17.0,4600.0],[48.6, 18.9,5000.0]], columns=X.columns)
# predict the species of the penguins on the ice flow using your best kNN model
...
```

## Wrap-up
Because randomness is involved for k-fold, I can't be sure exactly what your results were, but you definitely should not end up with all Gentoo!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 17

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for penguins (points: 1)
```
>>> assert len(penguins) > 0
```
