```python
# Initialize Otter
import otter
grader = otter.Notebook("hw1.ipynb")
```

## <u>Name:</u> [Your name] and (your onyen)



I encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 1: Introduction to Programming (35 points)

Please complete this notebook by filling in the cells provided. Be sure to run the cell above that loads the autograder, and be sure to save your notebook before generating your submission by running the last cell.  Also be sure to edit the cell above with your name and onyen!

!! Please submit a PDF of your notebook and upload it to Gradescope along with the generated zip file.  If no PDF submitted, **points will be deducted**.  You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop.  You can then drag and drop the zip & pdf file together to the upload box on Gradescope.

!! Points will also be **deducted** if you do not include your name and Onyen in the Markdown cell above!

Note that you are not expected to understand all the Python code in this homework; you will mostly be executing code and analyzing the output, but you will not be writing extensive amounts of code.

Also, note that this homework touches on many of the topics we will be covering in more depth during the course, so you are not expected to fully understand the concepts yet, but answer every question to the best of your ability.  Also, several of the questions do not have a single correct answer, so apply critical thinking when answering the questions.  We will discuss the homework problems and answers in class later.

**Deadline:**

This assignment is **due Tuesday, 9/2 at 8m ET**.

You should start early so that you have time to get help if you're stuck.

**Note: The tests included in this homework are meant to help with learning, but are not thorough grading tests.  That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be grading the homework manually.**



## 1. Data Types (5 points)



**Problem 1.1 (1 point).** Understanding the purposes and advantages of data types is an important aspect of being a proficient programmer.  Four basic types are `int`, `float`, `bool`, and `str` (integer, floating point, Boolean, and string).  As mentioned in class, all numeric data can also be represented as a string. So why not just make all data a string and just have one type? 

_Type your answer here, replacing this text._


**Problem 1.2 (4 points).** Using each exactly once, which basic data type (`int`, `float`, `bool`, or `str`) would work for the following four kinds of information related to a home listing?

**Situation 1.** Lot size in acres

_Type your answer here, replacing this text._

**Situation 2.** Street name and address

_Type your answer here, replacing this text._

**Situation 3.** Zip code

_Type your answer here, replacing this text._

**Situation 4.** Is a single-story dwelling

_Type your answer here, replacing this text._

## 2. Python: Assignment and Function calls (8 points)  
For each of the sentences below, write the corresponding Python code.  For example, for the sentence "Assign the string `Rick is left-handed` to the variable named `inconceivable`", the corresponding code would be:

`inconceivable = 'Rick is left-handed'`

Don't change any of the code that is provided, and **be sure to run the code cell below first!**

```python
# RUN THIS CELL BEFORE RUNNING OTHER CELLS IN PROBLEM 2
import pandas as pd
x = 5
y = 6.01
z = -4
percentage = 42.4818
name = 'Piper'
```

**Problem 2.1 (1 point).** Assign a value of 7.3 to a variable named `easypeasy`.

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q2.1")
```

**Problem 2.2 (1 point).** Assign the sum of `x` and `y` to a variable named `mysum`.

```python

# Write your code here by replacing the ...
...
```

```python
grader.check("q2.2")
```

**Problem 2.3 (1 point).** Assign the value of `x` raised to the fifth power to a variable named `x_to_the_fifth` 

```python
# Write your code here by replacing the ...
...
```

```python
grader.check("q2.3")
```

**Problem 2.4 (1 point).** Use code to concatenate the three strings `'Go'` and `' '` and `'Heels!'` and assign the result to a variable named `cheer_words` 

```python

# Write your code here by replacing the ...
...
```

```python
grader.check("q2.4")
```

**Problem 2.5 (1 point).** Using a built-in Python function, assign the data type of the variable `y` to the variable `detected_type`

```python

# Write your code here by replacing the ...
...
```

```python
grader.check("q2.5")
```

**Problem 2.6 (1 point).** Using a built-in Python function, assign the length of the string variable `name` to the variable `name_length`.

```python

# Write your code here by replacing the ...
...
```

```python
grader.check("q2.6")
```

**Problem 2.7 (1 point).** Using a built-in Python function, round the value of the variable `percentage` to the tenths place and assign the result to the variable `percentage_rounded`.

```python

# Write your code here by replacing the ...
...
```

```python
grader.check("q2.7")
```

**Problem 2.8 (1 point).** Using a built-in Python function, assign the minimum value of `x`, `y`, `z` to a variable named `lowest_coord` 

```python

# Write your code here by replacing the ...
...
```

```python
grader.check("q2.8")
```

## Problem 3. Python Lightning Round (10 points)
For each of the sentences below, write the corresponding Python code.  For example, for the sentence "Assign the string `Rick is left-handed` to the variable named `inconceivable`", the corresponding code would be:

`inconceivable = 'Rick is left-handed'`

Don't change any of the code that is provided.  Note that answers may require more than a single line of code.

**Problem 3.1 (1 point).** Assign 8.2 to two variables named `a` and `b` then assign the sum of `a` and `b` to a variable named `c`.

```python
...
```

**Problem 3.2 (1 point).** Assign a list of 4 consecutive even integers starting with 12 to a variable named `some_evens`.

```python
# Write your code here by replacing the ...
...
```

**Problem 3.3 (1 point).** To a variable named `my_items`, assign a list that contains -7, the string 'notebook', and 4.8 (in this order).

```python
...
```

**Problem 3.4 (1 point).** Assign a range of integers to a variable named `buzz` that starts at 7 and steps by 7 up to and including 35.

```python
...
```

```python
grader.check("q3.4")
```

**Problem 3.5 (1 points).** Assign the last item from a list named `many_things` to a variable named `last_one`.  Don't change any of the given code, only replace the three dots.

```python
import random
many_things = random.sample(range(0, 11), 10)
last_one = ...
```

**Problem 3.6 (1 points).** Assign to a variable named `third_one` the third item from a list named `more_things`

```python
import random
more_things = random.sample(range(100, 121), 10)
third_one = ...
```

**Problem 3.7 (2 points).** Write code to loop through all the elements in a list of strings named `lots_of_words` and increment a variable named `big_word_count` by a value of 1 for each string that is more than 4 letters long.  Remember, don't change any of the given code.

```python
lots_of_words = "This is an awful large number of words, don't-ya-know".split()
big_word_count = 0
...
```

**Problem 3.8 (2 points).** Write code to to loop through all the strings in a list of strings named `many_words` and print the string, the index of the string in the list, and the string converted to upper case.  Remember, don't change any of the given code.  This question will be graded manually.

```python
many_words = "So many words, such little oxygen in the room.".split()
...
```

## Problem 4. Control flow (12 points) 
In this problem you will be building on what you learned during Lab 3, so fully completing that is recommended.  You'll be writing code to loop through a list and perform different actions depending on the type of the element in the list, which is a common programming pattern.  This question will be manually graded.




**Problem 4.1 (2 points).** First, write code to loop through a list named `employee_data` and print the type of every item in the list.

```python
employee_data = ['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42]
# write your code here by replacing the "..." with your solution
...
```

**Problem 4.2 (2 points).** Next, rewrite your loop again below, but instead of printing the type of every one of the items in the list, test if the type of the item is a string, and if so, print the item.

```python
employee_data = ['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42]
# write your code here by replacing the "..." with your solution
...
```

**Problem 4.3 (2 points).** Next, copy your code from Problem 4.2 and add to it another test.  In the case that the type of the item was not `str`, test if the type is `int`, and if it is, add the item to `int_total` and print "Running total" followed by `int_total` on the same line.  You shouldn't need to change any of the code you've already written, just add more code.  Hint: the output when you run your completed code should look like this:
```
Rick
Running total 6
Piper
Running total 48
```

```python
employee_data = ['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42]
int_total = 0
# write your code here by replacing the "..." with your solution
...
```

**Problem 4.4 (1 point).** Copy your code from Problem 4.3, and add another test to it.  In the case the item is not a string or int, test if the item is a float and if so, increment `float_count` and add the item to `float_total` and print "Running average" followed by `float_total` divided by `float_count` on the same line.

```python
employee_data = ['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42]
int_total = 0
float_total = 0
float_count = 0
# write your code here by replacing the "..." with your solution
...
```

**Problem 4.5 (1 point).** Copy your from Problem 4.4, and this time add code that prints "Unprocessed item" for any type that you didn't already test for.  Hint: you don't need to add any new test for this.

```python
employee_data = ['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42]
int_total = 0
float_total = 0
float_count = 0
# write your code here by replacing the "..." with your solution
...
```

**Problem 4.6 (2 points).** Oops!  We actually are going to handle one more type.  Copy your code from Problem 4.5, and again you will just be adding to it.  Add a test to see if the item is type `range`, and if it is print "Counting..." then print every item in the range on a separate line.  Hint: the output should look like this:
````
Rick
Running total 6
Unprocessed item
Running average 8.6
Piper
Running average 16.45
Counting...
0
1
2
3
4
Running total 48
````

```python
employee_data = ['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42]
int_total = 0
float_total = 0
float_count = 0
# write your code here by replacing the "..." with your solution
...
```

**Problem 4.7 (2 points).** Okay!  One last bit of code to write.  Lists can contain objects of various types...including other lists!  For this problem, you'll be using all your code from Problem 4.6, but you'll be "nesting it" inside another `for` loop.  The "outside" loop will process a list named `company_data` that contains the data for multiple employees, with each employee's data in a list (just like you've been using).  The "inside" loop is your code from Problem 4.6.  One more thing (really?!?): use `enumerate` with the outer `for` loop and print "Employee #" followed by the index of the outer loop before executing the inner loop.

```python
company_data = [['Rick', 6, True, 8.6, 'Piper', 24.3, range(5), 42],
                ['Harlin', False, 'Puppycat', 9.3, range(6,0,-2)],
                ['Chudi', 1,2,3, 'No pets']]
# write the code for the outer loop here by replacing the "..." with your solution
...
    int_total = 0
    float_total = 0
    float_count = 0
# write the code for the inner loop here (your code from Problem 2.7)
...

```

Hint: the final output should look like this:
````
Employee # 0
Rick
Running total 6
Unprocessed item
Running average 8.6
Piper
Running average 16.45
Counting...
0
1
2
3
4
Running total 48
Employee # 1
Harlin
Unprocessed item
Puppycat
Running average 9.3
Counting...
6
4
2
Employee # 2
Chudi
Running total 1
Running total 3
Running total 6
No pets
````

You're all done with Homework 1!

**Important submission steps:**

1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**.
3. Click the link to download the zip file.
4. Go to Canvas and submit the zip file.
**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework1

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2.1 (points: 1)
```
>>> assert easypeasy == 7.3
```


Tests for q2.2 (points: 1)
```
>>> assert x == 5
>>> assert y == 6.01
>>> assert mysum == 11.01
```


Tests for q2.3 (points: 1)
```
>>> assert x == 5
```


Tests for q2.4 (points: 1)
```
>>> assert cheer_words == 'Go Heels!'
```


Tests for q2.5 (points: 1)
```
>>> assert y == 6.01
>>> assert detected_type == float
```


Tests for q2.6 (points: 1)
```
>>> assert name_length == len(name)
```


Tests for q2.7 (points: 1)
```
>>> assert percentage == 42.4818
>>> assert percentage_rounded == 42.5
```


Tests for q2.8 (points: 1)
```
>>> assert lowest_coord == -4
>>> assert lowest_coord == min(x, y, z)
```


Tests for q3.1 (points: 1)


Tests for q3.2 (points: 1)


Tests for q3.3 (points: 1)


Tests for q3.4 (points: 1)
```
>>> assert type(buzz) == range
```


Tests for q3.5 (points: 1)


Tests for q3.6 (points: 1)


Tests for q3.7 (points: 2)
