```python
# Initialize Otter
import otter
grader = otter.Notebook("lab11.ipynb")
```

# Lab 11: Implementing Randomness


**Submission instruction**: Please upload your `lab11.ipynb` file and a `pdf` to Gradescope. This lab will be manually graded (no autograder).


```python
# write your code here
me = ["Rico Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

```python
grader.check("yourname")
```

## 1. Generating random data (5 minutes)
Suppose you are programming an RPG (Role Playing Game) and your game follows the common practice of considering it a "critical hit" when you roll a 20 on a 20-sided die.


### Using the random package
Implementing randomness on a machine that is designed to be completely deterministic is actually quite challenging.  Luckily, there is a Python package for that!  Copy the code below to import the `random` package.
```Python
import random
```



```python
# copy code here
...
```

Now, implement a function named `roll_d20()` that returns a random integer from 1 to 20 inclusive. To do this, the best function is `randint` from the `random` package.  Implement your function below.

```python
# write your function here
...
```

### Testing
Like good software developers always do, test your function!  In the cell below, write code to call your function and print the result, and run the cell several times until you get a 20 and 1 at least once.  If you never get a 20 or a 1, or get a number that is not between 1 and 20, maybe you have a bug....  

```python
# test your function here
...
```

Doing testing by calling the function once at a time is rather tedious.  Surely you could write some Python code that would, say, call your function 20 times....

```python
# write your code here
...
```

## 2. Is this random? (10 minutes)
You integrate your random function into the game, and the QA department starts testing.  Part of their job is to record every die roll they get while playing the game.  (Yes, game testers have to do more than just play games all day.  Talk about tedious!)

Your company employs 3 game testers Vivian, Zoe, and Mitchell.  There has been some chatter that one of the testers is just playing for fun, then fabricating their die roll data later, so you are called into investigate.  Copy the code below to load in the testing data:
```Python
import pandas as pd
tester_data=pd.read_csv("tester_data.csv")
```

```python
...
```

It is always a good idea to check out your data to see what you are working with.  Things like `head()`, `tail()`,`shape` are a great way to start. Give those 3 a try.

```python
# write your code here
...
```

```python
# write your code here
...
```

```python
# write your code here
...
```

Yikes!  10,000 die rolls per tester!  That is a lot!  And a lot of data for you to analyze!  Luckily, you took DATA110H so you have a plan....
###  Investigation 1: Compute summary statistics (1 minute)
First, you decide to check out the summary statistics of the data.  This is super easy using `describe`.  Do this below.

```python
# write your code here
...
```

What do you notice?  Anything fishy yet?  Hmmm, not really.
### Investigation 2: Look at the distribution (5 minutes)
What about doing some histplots?  Do an individual histplot for each of the testers.  How many bins do you think it makes sense to use?  You can either do this in separate cells, or get fancy and do it in one cell as 3 subplots.

```python
# write your code here
...
```

Anything interesting?  Describe below.

// Write your answer here.

### Investigation 3: Looking for runs
You vaguely remember that one lab in DATA110H when professor Marks amazed you with his ability to detect fake data, until you learned the "trick" was to look for runs, or the absence of runs.  You decide to check for runs of 3 in the data.  But first, these names are starting to make you uncomfortable, so you decide to anonymize the data a bit and convert each column of data to a separate list.  Copy this code to do that:
```Python
tester1 = list(tester_data.Vivian)
tester2 = list(tester_data.Zoe)
tester3 = list(tester_data.Mitchell)
```

```python
# copy code here
...
```

First, you decide to warm up by just writing a function to count the number times 20 appears in the data.  Write a function named `count_critical()` that takes a list as an argument and returns the number of times 20 appears.  There are many ways to do this with Python, but humor me and use a `for` loop this time.  And to make it easier to grade, please name the argument to the function `dierolls`


Now, write a function named `count_critical_run_of_3()` that takes list as input and returns the number of times 20 appears in the data 3 times in a row. 

```python
# write your function here
...
```

Let's see the results: copy the code below to print out the number of criticals:
```Python
print(count_critical(tester1), count_critical(tester2), count_critical(tester3)) 
```

```python
# copy code here
...
```

You can verify the results by looking at your histplots, or by using the pandas `value_counts` method on each table column.

Okay, now you are all warmed up.  Write a function named `count_critical_run_of_3()` that takes list as input and returns the number of times that 20 appears in the data 3 times in a row.  Hint: start with the code you wrote for `count_critical` and modify it.

```python
# write your function here
...
```

Run your function for each of the 3 testers with this code:
```Python
print( count_critical_run_of_3(tester1), count_critical_run_of_3(tester2), count_critical_run_of_3(tester3) )
```

```python
# copy code here
...
```

Yes, things are starting to become clear!  Maybe just one more investigation....
### Investigation 4: Looking for gaps
Similar to looking for runs, another thing you can check for is gaps; the longest streak of numbers without a 20 being rolled.  This is a little trickier, so here is some code to get started:
```Python
def find_biggest_gap(dierolls):
    biggest_gap = 0
    current_gap = 0
    for roll in dierolls:
        if roll == 20:
            ...
        else:
            ...
    return biggest_gap
```
Copy this code below and replace the `...` sections with the necessary logic.

```python
# write function here
...
```

Check for gaps in the data with this code:
```Python
print(find_biggest_gap(tester1), find_biggest_gap(tester2), find_biggest_gap(tester3) )
```

```python
# copy code here
...
```

Yes, one of the testers is definitely looking suspicious.  Which one is it (1,2,3)?  No names please!  

// Write your answer here

# BONUS
The designer of your game happens to be walking by while you are doing your analysis, and sees the run data you calculated.  They make the snap decision that the game would be better if players never get 2 critical hits in a row.  So they ask you to modify your `roll_d20` function so that it doesn't return a 20 if the previous roll was a 20.  Implement this new version:
```Python
def roll_d20_no_critical_repeat(prev_roll):
   ...
```
and then test it.

```python
# write function here
...
```

```python
# write test code here
...
```

# BONUS BONUS
What a horrible idea!  Turns out everyone loves getting critical hits (go figure!).  So instead, they ask you to actually make it so that the chance of getting a critical hit increases by 5% for every roll you don't get one.  So right after a critical hit, your chance is 1 in 20, or 5% (the standard).  Then if you didn't get a critical hit, your next roll will have a 10% chance.  If that one isn't critical, then your next roll will have 15%, etc.  Implement the new function with the desired behavior:
```Python
def roll_d20_adjusted(rolls_since_20):
    ...
```
Use this code to test your function:
```Python
rolls_since_20 = 0
for i in range(25):
    roll = roll_d20_adjusted(rolls_since_20)
    print(roll)
    if roll == 20:
        rolls_since_20 = 0
    else:
        rolls_since_20 += 1
```

```python
# write function here
...
```

```python
# copy test code here
...
```

Does it seem like you are getting more criticals?  What is the longest possible gap between criticals now?  What is your chance of getting a 1 on the first roll?  If you didn't get a 20 on the first roll, what is your chance of getting a 1 on the second roll? (that one is trickier, but you should be able to do it)

// write you answers here

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 8

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for yourname (points: 1)
```
>>> assert me[0] != 'Rick Marks'
```
