```python
# Initialize Otter
import otter
grader = otter.Notebook("hw7.ipynb")
```

## <u>Name:</u> [Your name] and (your onyen)


The final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.




Please complete this notebook by filling in the cells provided.

!! Please submit a  zip and **PDF** of your notebook. You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop. All problems will be manually graded. **Make sure all your code is visible and not cut off in the pdf. If we can't see your answer code for a coding problem, points may be deducted.**

!! Points will also be **deducted** if you do not include your name and Onyen in the Markdown cell above!!

# Homework 7: Hypothesis Testing (25 points)

In this homework, you'll be analyzing the fairness of a 6-sided die.

A street vendor asks you if you'd like to play a game.  The game is:
* pay \$10 to roll a 6-sided die 3 times
* if any of those rolls is a 2 he pays you $12.  

You decide to play the game 30 times, and out of your 90 total rolls of the die, you roll a 2 **seven times**.  You start to suspect something is not right, and you wonder if perhaps the die is NOT fair.  

## Problem 1. Loaded question (14 points)



#### **Problem 1.1 (4 points) Hypothesis Testing.**
In particular, you feel that there might be **less than fair chance** of rolling a 2, even though the street vendor insists it is fair. You decide to conduct a hypothesis test.

1) State the null hypothesis for this situation.  Remember, you'll need to be able to "simulate under the null".
2) State the alternative hypothesis.
3) Which of the following statistics would you be appropriate for evaluating these hypotheses?
* a) cost to play the game
* b) number of 5's rolled
* c) proportion of 2's rolled
* d) total number of rolls
4) What is the value of this statistic for the 30 games you played?  Assign it to `measured_statistic1` in the provided cell.

*Write your answers to 1-4 here*

```python
# Assign the value of the statistic for the 30 games here
measured_statistic1 = ...
```

#### **Problem 1.2 (10 points) Simulating under the null**
In this problem, you'll be writing code to simulate the test statistic for 90 rolls, and you'll run that simulation 1000 times. 

**Problem 1.2.1 (2 points).** 
First, write a function named `simulate_six_sixed_die` that takes a number of rolls as an argument and returns a list containing that many simulated 6-sided die rolls. Import whatever Python package you need to do this.

```python
# write your code here
...
```

```python
grader.check("q1.2.1")
```

**Problem 1.2.2 (2 points).**
Next, write code to simulate 1000 times the test statistic for 90 die rolls, and store the test statistic in a table named `statistic1`

```python
import pandas as pd
statistic1 = pd.DataFrame(columns=['test statistic'])
# finish the code below, replacing all the ... with your code
for simulation in range(1000):
    # simulate 90 rolls of a six-sided die using your function
    roll_list = ...
    simulated_rolls = pd.DataFrame(roll_list, columns=['roll']) # create a 1-column table of randomly simulated rolls
    # twos is an array of True or False. True if roll is 2, False otherwise.
    twos = ...
    # calculate the statistic for this simulated data using twos
    sim_t1 = ...
    statistic1.loc[len(statistic1)] = [sim_t1] # add simulated result to end of table
statistic1
```

**Problem 1.2.3 (2 points).**
Plot the empirical distribution of your statistics simulated under the null.  Plot a vertical line for the value of the measured statistic for your actual games.  Label your plot using the parameter `label=` just like in the lec15 demo code.

```python
# 3. Plot empirical distribution of t1 (under the null)
import seaborn as sns
import matplotlib.pyplot as plt
fig, ax = plt.subplots(1)
# write your code here
...
ax.legend()
```

**Problem 1.2.4 (1 point).**
Looking at the shape of your distribution, what type of distribution is your simulation data, approximately? We're referring to the distribution of the statistics simulated under the null.

*Write your answer here*

**Problem 1.2.5 (2 points).**
Write code to compute the p-value from the `statistic1` table.

```python
p_value1 = ...
p_value1
```

**Problem 1.2.6 (1 point).**
Does your simulation provide evidence against the null?  (yes or no)  If so, is the evidence statistically significant?  Is it extremely statistically significant?  Explain the convention you are using to answer these questions.

*Write your answers here*

## Problem 2. (11 points) Reloaded question.
#### **Problem 2.1 (4 points) Hypothesis testing**
Instead of investigating if there is a **less than fair** chance of rolling a 2 like you just did in Problem 1, now evaluate if there is an **unfair chance** of rolling a 2 (either more likely or less likely than fair).

1) State the null hypothesis for this situation.
2) State the alternative hypothesis.
3) Which of the following statistics would you be appropriate for evaluating these hypotheses?
* a) number of 2's rolled in a row
* b) abs(proportion of 2's rolled - 1/6)
* c) abs(proportion of 2's rolled + 1/6)
* d) abs(total number of rolls)
4) What is the value of this statistic for the 30 games you played in Problem 1?  Assign it to `measured_statistic2` in the provided cell.

*Write your answers to 1-4 here*

```python
# Assign your measured statistic here
measured_statistic2 = ...
```

#### **Problem 2.2 (7 points) Simulating under the null revisited**

**Problem 2.2.1 (2 points).**
Now, write code to simulate 1000 times this new test statistic for 90 die rolls, and store the results in a table named `statistic2`.  You can use the same `sample_die()` function you wrote before.

```python
statistic2 = pd.DataFrame(columns=['test statistic'])
# finish the code below, replacing all the ... with your code
for simulation in range(1000):
    # simulate 90 rolls of a six-sided die using your function
    roll_list = ...
    simulated_rolls = pd.DataFrame(roll_list, columns=['roll']) # create a 1-column table of randomly simulated rolls
    twos = ...
    sim_t2 = ...
    statistic2.loc[len(statistic2)] = [sim_t2] # add simulated result to end of table
statistic2
```

**Problem 2.2.2 (2 points).**
Plot the empirical distribution of your statistics simulated under the null.  Plot a vertical line for the value of the measured statistic for your actual games.  Label your plot using the parameter `label=`.

```python
fig, ax = plt.subplots(1)
# write your code here
...
ax.legend()
```

**Problem 2.2.3 (2 points).** Write code to compute the p-value from the `statistic2` table.


```python
p_value2 = ...
p_value2
```

**Problem 2.2.4 (1 point).**
Does your simulation provide evidence against the null?  (yes or no)  If so, is the evidence statistically significant?  Is it extremely statistically significant?  

*Write your answers here*

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework7

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.2.1 (points: 2)
```
>>> assert type(simulate_six_sided_die(5)) == list
>>> assert len(simulate_six_sided_die(5)) == 5
>>> assert min(simulate_six_sided_die(1000)) == 1
>>> assert max(simulate_six_sided_die(1000)) == 6
```
