```python
# Initialize Otter
import otter
grader = otter.Notebook("lab8.ipynb")
```

# Lab 8: Matplotlib


**Submission instruction**: Please create a `zip` and `pdf` via File -> Print (or cmd + P on mac), and upload it to Gradescope. Most of the lab will be manually graded.


```python
# write your code here
me = ["Rick Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

```python
grader.check("yourname")
```

The last 2 weeks in lectures, we discussed how to use `seaborn` to visualize table data. As mentioned in the beginning, `seaborn` is a powerful tool that can get us very beautiful plots easily, but it can also be difficut to customize plots to 100% our liking. That's why today we are going over some important `matplotlib` functions and parameters that can be used in conjunction with `seaborn`.

```python
# Run this cell

import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
```

# 1. Figures and axes (5 minutes)

These two images from the [ListenData matplotlib tutorial](https://www.listendata.com/2019/06/matplotlib-tutorial-learn-plot-python.html) summarize the most important elements of a Matplotlib figure (remember that seaborn uses matplotlib behind the scenes).


![Matplotlib figure elements](https://blogger.googleusercontent.com/img/b/R29vZ2xl/AVvXsEj-ZStrYVLlanGuxuDmk_m2W8oXCIH-ueK9cgORQC_oY9_YwAkxf4npj2GbupXFaroy0s05fc2naWUsoF8J6R-A1pzKvsIKNn_ZmfYQqHKsVfmc0gVM2jcH6ZCb4ykO0H847mWyCYctwy0C/s1600/basics_matplotlib.PNG)

- A **figure** is the overall space in which visualizations may be created. Think of it as a canvas for your painting.
- An **axis** is the x- or y- axis along which the values occur.
    - It may have a text **label** that describes what type of values are occurring along that direction.
    - It may have **ticks** that mark spots on the axis. These ticks may have corresponding **tick labels**.
- An **axes** is a single visualization (i.e. "a plot") containing both x- and y-axis.
    - It may contain a **title**.
    - Figures can contain one or more axes (just like you can have multiple drawings and paintings on one canvas). See the image below.

![Subplots](https://blogger.googleusercontent.com/img/b/R29vZ2xl/AVvXsEhUh6lQj_vTkc0aoAnNVLuD-mmFNmn24dCj9hoPtdeABBnw4gPR0DNW-sYGLDQhIKJiTgU7bIbaAfMPXVXCF1qa6bBKxNGpMKKSc-cvAASK7Dk9wKFvQta2qn3G_5HLi9myL4LbFHQDqTEd/s1600/figure_axes.PNG)



Here is how you make a figure and *an* axes.

```Python
fig, ax = plt.subplots(1)
```

```python
# Copy and run the code here

```

Specify the figure dimension (width, height) in inches with parameter `figsize`.

```Python
fig, ax = plt.subplots(1, figsize=(5,3))
fig, ax = plt.subplots(1, figsize=(12,8))
```

**Q**: Do you notice anything different about the size of the box as well as the fontsizes? 

```python
# Copy and run the code here

```

# 2. Combining axes with seaborn (15 minutes)

This is actually fairly simple.
1. Create a matplotlib axes. This is equivalent to telling seaborn that this is the canvas we're now working with.
2. Run seaborn plotting functions.
3. Update parameters of the axes.

Here is an example that creates a figure only for the purpose of demonstration:

```Python
fig, ax = plt.subplots(1, figsize=(8,8))
sns.scatterplot(data=penguins, 
                x='Body Mass (g)', 
                y='Flipper Length (mm)',
               color = '#056412',
               marker= 'x')

ax.set_title('Palmer Penguins', fontsize=20)
ax.set_ylabel('Flipper Length (cm)', fontsize=16)
ax.set_xlabel('Body Mass (kg)', fontsize=16)

ax.set_xticks([3000, 4000, 5000, 6000])
ax.set_xticklabels([3,4,5,6])

ax.set_yticks([160, 180, 200, 220, 240])
ax.set_yticklabels([16, 18, 20, 22, 24])
# END SOLUTION
```

Compare it with the default plot you get with
```Python
sns.scatterplot(data=penguins, 
                x='Body Mass (g)', 
                y='Flipper Length (mm)')
```

**Q**: Can you figure out what each of the added code does?
- In particular, what does `alpha` do? Hint: It must be a float between 0 and 1.

**Hint**: 
Helpful links:
- [List of parameters](https://seaborn.pydata.org/generated/seaborn.scatterplot.html) accepted by seaborn scatterplot
- [Additional list of matplotlib parameters](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.scatter.html#matplotlib.axes.Axes.scatter) that are also accepted by seaborn scatterplot
- [List of markers](https://matplotlib.org/stable/api/markers_api.html#module-matplotlib.markers)


**Remark**: An important note on colors:
- There are [many ways to specify colors](https://matplotlib.org/stable/users/explain/colors/colors.html#color-formats) for matplotlib. For example, you can provide a hex color code such as `'#056412'`.
- When choosing colors, it is adivsable to use color palettes that are accessible to those with color blindness. See resources on this topic such as [this](https://www.nceas.ucsb.edu/sites/default/files/2022-06/Colorblind%20Safe%20Color%20Schemes.pdf) and [this](https://davidmathlogic.com/colorblind/).


```python
# run this cell

penguins = pd.read_csv('~/110H-F25/demo/palmer_penguins.csv')
penguins = penguins.dropna(subset=['Culmen Length (mm)', 'Culmen Depth (mm)'])
```

```python
# Copy and run the code here

```

```python
# Copy and run the code here

```

# 3: Saving your figure (5 minutes)

Last step is saving your plot as a file! This is where you get to use `fig`.

```Python
fig.savefig('penguins-scatter.png') 
```

Whether you want `jpg` vs `png` images depends on the situation. `png` contains more information than `jpg` and therefore can be considered higher quality, but `jpg` files are smaller and will therefore load faster on a website.

No need to submit this file to Gradescope.

```python
# Copy and run the code here


```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 8

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for yourname (points: 1)
```
>>> assert me[0] != 'Rick Marks'
```
