```python
# Initialize Otter
import otter
grader = otter.Notebook("lab14.ipynb")
```


**Submission instruction**: Please create a zip and `pdf` via File -> Print (or cmd + P on mac), and upload it to Gradescope. No autograder since there are many correct ways to explore the data. You are **not** required to finish the entire worksheet provided that you were actively engaged during the entire class period. As long as you've collaborated, put in good effort and made **reasonable progress** during the lab period, you can expect to get full credit. So even if you decide to finish it at home, please submit what you have by the end of 50 minutes to make sure you get credit! 

#  Lab 14: Data Exploration for Final Project Proposal

Your project proposal must contain all required sections outlined in the project document.

During this lab, your will work on the data exploration part of the project proposal. By the end of the lab, you must submit progress on **data exploration** ***individually***. 


```python
# edit these names to your name and your final project partner's name
me = ["Rick Marks"]
partner = ["Piper Marks"]
...
```

```python
grader.check("names")
```

## Part 1: Work with your final project teammate to decide on your dataset (15 mins)


Specify the dataset your project intends to use.  Datasets and info for some topics can be found in the DATA110H/project/ directory.  For others, you will need to get the data yourselves from the web (see project page at https://docs.google.com/document/d/1reaiF54ByxpmVjM4n35J4PREyJ82CpWFtUoWcEqLnCQ/)

```python

```

## Part 2: Data exploration individually (20 mins)

Conduct data exploration: load your data into Python and calculate summary statistics and create **at least two basic visualizations**.  Add as many code cells as you need.

Try to create visualizations that are different from those created by your teammates.

```python

```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 14

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for names (points: 2)
```
>>> assert me[0] != 'Rick Marks'
>>> assert partner[0] != 'Piper Marks'
```
