```python
# Initialize Otter
import otter
grader = otter.Notebook("lab4.ipynb")
```

# Lab 4. Wordle

In this activity, we will start using Pandas, though most of the lab will really be about writing basic Python to implement a playable version of the game Wordle!

```python
# write your code here
me = ["Ric Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

```python
grader.check("q0")
```

## 1. Loading a table from a file, and understanding the table. (5 minutes)
First, you will load some Comma-Separated Values (CSV) data from a file.  You'll use the pandas function `read_csv()` to do this, so pandas needs to be imported (which is done for you in the code below).  

**Q**: Add one line of code that calls `pd.read_csv()` to read a file named 'wordle_answers.csv', and assign the table to a new variable `wordle_df`. You can do this by replacing the ... with your code.

```python
import pandas as pd
wordle_df = ...
```

```python
grader.check("read")
```

**Remark**: Run the following cell to learn more about the data type of the Pandas table.

```python
type(wordle_df)
```

Note that sometimes when you check the type of something, the return value looks more complex than what you would expect, but it is still equivalent to the simpler type that you were expecting, as you can see in the next cell.

```python
type(wordle_df) == pd.DataFrame
```

**Q**: Run each of the following cells. Make sure to understand what kind of information they give you. Be careful - some of them are calling functions (table methods ), so require parenthesis `()`, while some are just looking up information attached to the table (think of it like metadata in your phone images), so they don't involve `()`.

```python
wordle_df.dtypes
```

```python
wordle_df.columns
```

```python
wordle_df.shape
```

```python
wordle_df.info()
```

```python
wordle_df.describe()
```

## 2. Adding columns (5 minutes)

**Q**: Write one line of code below that selects the column `word` in table `wordle_df`. We showed you two ways to achieve this in lecture.

```python
# Write your code here!


```

We want to keep track of whether a word has been used in a game or not so we don't repeat words.  One way to do this is to add a `bool` attribute for each entry.  As you know, attributes are columns in tables, so you'll need to add a new column.

Add the column `already_used?` to `wordle_df` and initialize all the values to `False` 

```python
...
wordle_df
```

```python
grader.check("column")
```

The classic Wordle game has all letters upper case, but you may have noticed all the letters in our table are lower case.  Add an attribute named `word_upper` to `wordle_df` that is the upper case version of the words. 
To do this, use one of the special `str` column operations that converts all letters to upper case.

```python
...
wordle_df
```

```python
grader.check("column2")
```

## 3. Let's design Wordle!
Implementing Wordle is a great way to hone your basic Python skills.  But just like with any nontrivial programming problem, it is important to first think about the problem and understand it before you jump into coding.  In case you have never played Wordle:
https://www.nytimes.com/games/wordle/index.html

We'll be using the same color scheme for giving feedback as classic Wordle: if a guessed letter is printed green, it means it is in the answer and in the correct spot.  If a guessed letter is printed yellow it means it is means the letter is in the answer but not that spot.  If a letter is black, it is not in the answer.  Don't worry, I will provide you the code to colorize what you will print!

Before continuing, stop and discuss with your lab partner some of the steps you think would be key to implementing Wordle.





*feel free to write your steps here*
1.
2.
3.
4.
5.
6.

Here is my recommended breakdown:

#### Wordle steps
1. Randomly get an answer word from the table 
2. Give the player 6 guesses at the answer word
3. __Verify each guess is valid (5 letters long, and a potential answer word)
4. __Evaluate each valid guess against the answer word and give appropriate colored feedback.
5. __If a valid guess matches the answer word, no more guesses are needed.  Stop and print a success message.
6. If the player doesn't guess the word after 6 tries, print the answer word.

Note the not very subtle indentation for steps 3, 4, 5.  Why do you think I did this?

## 4. Let's implement Wordle!
The below code implements the steps above, with some more parts to be added (TBD).  Take some time to make sure you understand this code before you move on to implement the TBD sections.

````
# STEP 1
import random
answer = random.choice(wordle_df['word']) # get a random word from wordle_df

# STEP 2
correct_guess = False
for attempt_num in range(1,7):

    # STEP 3: Get a valid guess (TBD)
    guess = "temp" # placeholder code

    # STEP 4: Evaluate the guess against the answer and get the feedback string (TBD)
    feedback_string = "*****" # placeholder code
    
    print(f"Guess #{attempt_num}: {guess}, Result: {feedback_string}")

    # STEP 5: if the guess matches the answer exit the loop
    if guess == answer:
        correct_guess = True
        break
# STEP 6: print the end-of-game feedback
if correct_guess==False:
    print("Better luck next time.  The correct word was:", answer)
else:
    print("Yay! You guessed the word!")
````

### STEP 3: Valid user input
Your first job is to replace the placeholder code above for STEP 3.  This will be a lot like what you did in Lab3 for the password, but this time you should use a `while` loop.  There are while loop examples in `demo/lec4.ipynb` that may help.  To validate the guess, you should:
1. ask the player to input a guess (like Lab3)
2. convert the guess to uppercase letters
3. check if the guess is 5 characters long and all the characters are alphabetical characters.
4. check if the guess is a possible answer from `wordle_df`
5. do all this again if the guess is not valid

Hint: the string methods `upper()` and `isalpha()` and the built-in function `len` might be useful! 



```python
# Write your code here!

```

There are multiple ways you can do this, but if you are stuck, the following rough outline is a good starting skeleton:
````
is_valid_guess = False
while ...:
    guess = ...
    guess = guess.upper() # convert guess to uppercase
    if ...:
        print("Invalid guess. Please enter a 5-letter word.")
    elif ...:
        print("Word not in answer list. Please enter a valid 5-letter word.")
    else:
        is_valid_guess = True
````

### STEP 4: Evaluating the guess against the answer
Evaluating the guess against the answer can be done in many different ways, but we are going to break it into 2 steps:
1. Find what letters of `guess` are in `answer` and in the right place (the green letters)
2. Find what letters of `guess` are in `answer` but in the wrong place (the yellow letters)
3. Create color-coded feedback string

#### STEP 4.1. Right letter right place
The first step is pretty straight forward to implement, since at this point in the code we've made sure that `guess` and `answer` are both 5-letter uppercase words.  Use a list of 5 bool values named `right_letter_right_location` to hold the results.  Fill in the missing code using `guess` and `answer` and string indexing/comparison:
````
right_letter_right_place = [False, False, False, False, False]
for i in range(...):
    if ...:
        right_letter_right_place[i] = True
````

```python
# Write your code here!
# You can test this code by temporarily setting values for guess and answer and printing right_letter_right_location
# For example:
# guess = "CRANE"
# answer = "CRATE"


```

#### STEP 4.2 Right letter wrong place
The second step is a bit trickier logic, though the code is not very long.  But too tricky to figure out in the few minutes of lab!: 
````
right_letter_wrong_place = [False, False, False, False, False]
answer_letter_list = list(answer)
for i in range(5):
    if right_letter_right_place[i] == True:
        answer_letter_list[i]=' ' # mark as used
for i in range(5):
    if right_letter_right_place[i] == False:
        if guess[i] in answer_letter_list:
            right_letter_wrong_place[i] = True
            location_of_guess_letter_in_answer = answer_letter_list.index(guess[i])
            answer_letter_list[location_of_guess_letter_in_answer] = ' ' # mark as used
````

#### STEP 4.3 Color-coded feedback string
The code to colorize the feedback string uses "escape codes":
````
feedback_string = ""
for i in range(5):
    if right_letter_right_place[i] == True:
        feedback_string += "\033[42m" # GREEN
    elif right_letter_wrong_place[i] == True:
        feedback_string += "\033[43m" # YELLOW
    feedback_string += " " + guess[i] + " "
    feedback_string += "\033[00m" + " " # add a space between letters
````

## 5. Putting it all together
Okay, so maybe this was a bit ambitious!  Here is a working version of Wordle complete with all the steps copied over:

````
# STEP 1
import random
answer = random.choice(wordle_df['word_upper'])

# STEP 2
correct_guess = False
for attempt_num in range(1,7):
    # STEP 3
    is_valid_guess = False
    while is_valid_guess==False:
        guess = input("Enter your Wordle guess: ")
        guess = guess.upper()
        if len(guess) != 5 or guess.isalpha() == False:
            print("Invalid guess. Please enter a 5-letter word.")
        elif guess not in list(wordle_df['word_upper']):
            print("Word not in list. Please enter a valid 5-letter word.")
        else:
            is_valid_guess = True

    # STEP 4
    # STEP 4.1
    right_letter_right_place = [False, False, False, False, False]
    for i in range(5):
        if guess[i] == answer[i]:
            right_letter_right_place[i] = True

    # STEP 4.2
    right_letter_wrong_place = [False, False, False, False, False]
    answer_letter_list = list(answer)
    for i in range(5):
        if right_letter_right_place[i] == True:
            answer_letter_list[i]=' ' # mark as used
    for i in range(5):
        if right_letter_right_place[i] == False:
            if guess[i] in answer_letter_list:
                right_letter_wrong_place[i] = True
                location_of_guess_letter_in_answer = answer_letter_list.index(guess[i])
                answer_letter_list[location_of_guess_letter_in_answer] = ' ' # mark as used

    # STEP 4.3
    feedback_string = ""
    for i in range(5):
        if right_letter_right_place[i] == True:
            feedback_string += "\033[42m" # GREEN
        elif right_letter_wrong_place[i] == True:
            feedback_string += "\033[43m" # YELLOW
        feedback_string += " " + guess[i] + " "
        feedback_string += "\033[00m" + " " # add a space between letters
    print(f"Guess #{attempt_num}: {feedback_string}")

    # STEP 5
    if guess == answer:
        correct_guess = True
        break
# STEP 6
if correct_guess==False:
    print("Better luck next time.  The correct word was:", answer)
else:
    print("Yay! You guessed the word!")
````

```python
# copy code here to try Wordle
```

## BONUS
The real Wordle doesn't print "Yay! You guessed the word!" when you get it right.  It prints a different message depending on hom many attempts it took you to get the word:

1: Genius 2: Magnificent 3: Impressive 4: Splendid 5: Great 6: Phew!

Try changing the implementation to remove the "Yay!" print and give this feedback instead using multiple conditional tests and prints (using multiple `elif` commands, and possible `if` or `else` depending on you choose to do it). 

```python
# write your code here!
```

Now try doing the feedback differently by sticking all the feedback message strings into a list, and do a single print that prints one item from the list using indexing. 

```python
# write your code here!
```

## BONUS
Update the implementation to verify the value of `aleady_used?` for the answer word is `False`.  If it isn't, you should get another random answer word from the table.  What if the next random word is already used too?  Be sure to handle that case!  Hint: the new Python commands from Lecture 4 seem perfect for this. 

```python
# Write your code here!
```

## BONUS BONUS
People look at all sorts of trivia related to Wordle words.  Add an attribute to `wordle_df` that is the number of vowels **(a, e, i, o, u, y)** in each potential Wordle answer. 
The attribute should be named `vowel_count` and should be initialized with the number of vowels in each word.  To calculate that, you can use a special `str` column operation that counts how many times a character appears in a string. For example, to count the letter `x` for every string in the column `wordle_df['name']`:

`wordle_df['name'].str.count('x')`



```python
# Write your code here!
```

But y is actually only **sometimes** a vowel.  A simple rule that is often used is that y is not a vowel when followed by a vowel (aeiou).  Try implementing this challenging twist!  

```python
# Write your code here!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 4

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for column (points: 2)
```
>>> assert 'already_used?' in wordle_df.columns
>>> assert wordle_df['already_used?'][0] == False
```


Tests for column2 (points: 2)
```
>>> assert wordle_df['word_upper'][0] == 'ABACK'
>>> assert wordle_df['word_upper'][1] == 'ABASE'
```


Tests for q0 (points: 1)
```
>>> assert me[0] != 'Rick Marks'
```


Tests for read (points: 1)
```
>>> assert wordle_df.loc[0, 'word'] == 'aback'
```
