```python
# Initialize Otter
import otter
grader = otter.Notebook("lab13.ipynb")
```

# Your name and your onyen


**Submission instruction**: Please save then create both a zip and `pdf` (via File -> Print or cmd + P on mac), and upload it to Gradescope.

#  Lab 13:  Sampling and Hypothesis Testing

<img src="Trick-or-Treat-Yourself.jpg" width=500px>
<div style="text-align: center"> </div>

## Part 1: Data Collection (15 mins)
We are going to be collecting data about **the number and kind of candies** that are taken when **a person "grabs a reasonable handful"**.  Every student reach into the bowl of candy with one hand, grab a handful of candy, and announce the number of **Starbursts**, number of **Kisses**, and number of **Miniatures** for the other students to record.  To be considered a valid "handful" sample, at least 1 candy must be pulled out of the bowl.  

**Questions: (4 points)** 
1. What are your initial thoughts/concerns about the plan?  
2. What is the population that is being sampled?
3. What kind of sampling do you think this would be considered?
4. Do you expect to get a representative sample?  

*Write your answers here*  

What is your guess for the average number of candies in a handful? **(1 point)**


```python
my_guess = ...
```

```python
grader.check("q1")
```

What kind of distribution do you think you'll see in the data you collect?  Why? **(1 point)**

*Write your answer here*  

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import random
```

```python
# USE THIS CODE FOR LOADING COLLECTING DATA
candy_df = pd.DataFrame(columns=['Starbursts','Kisses','Miniatures'], data=[
    # replace these placeholder measurements and add actual measurements in the form [?,?,?] (just like the placeholders)
    [1, 2, 3], [42,42,42],
...

    ])
candy_df
```

## Part 2 Exploratory Data Analysis (EDA) (5 minutes)
Run the code cell below.

```python
fig, ax = plt.subplots(1)
ax.set_ylabel('Kisses, Miniatures per Handful')
ax.set_xlabel('Starbursts per Handful')
sns.scatterplot(candy_df, x='Starbursts', y='Kisses',label='Kisses')
sns.scatterplot(candy_df, x='Starbursts', y='Miniatures',label='Miniatures')
ax.set_title('Scatter plot of Kisses and Miniatures vs. Starbursts per Handful')
```

From the plot above, does the # of Kisses or # of Miniatures correlate better with the # of Starburst per handful? **(1 point)** 

*Write your answer here*  

Add the code to compute the correlation table between the types of candies. **(1 point)**

```python
correlation_table = ...
correlation_table
```

Was your interpretation of the scatterplot correct?

Now write code to add another column to `candy_df` named `'Total'` that is the total number of pieces of candy in each sample. **(1 point)**

```python
# write your code here, replacing ...
...
candy_df
```

Next, let's examine that total.  First, compute the mean and standard deviation for your `Total` column and assign them to variables `handful_mean` and `handful_std` **(1 point)**

```python
handful_mean = ...
handful_std = ...
```

In the interest of time, below I've supplied code for several different visualizations of the Total # per handful.  Execute each of the cells to see the visualization; if they don't execute, you must have an issue with your code cells above.

```python
fig, ax = plt.subplots(1)
sns.barplot(candy_df['Total'],label='Candies per handful')
ax.axhline(handful_mean, color='g', label='Mean')
ax.set_xlabel('Student (one handful per student)')
ax.set_ylabel('Total # of candies in a handful')
ax.legend()
ax.set_title('Bar chart of Total Candies per Handful')
```

```python
fig, ax = plt.subplots(1)
sns.histplot(candy_df['Total'], label='Candies per handful', stat='density')
ax.set_xlabel('Total # of candies in a handful')
ax.set_ylabel('Density')
ax.axvline(handful_mean, color='g', label='Mean')
ax.axvline(handful_mean-handful_std, color='r', label='Mean +/- 1 std', linestyle=':')
ax.axvline(handful_mean+handful_std, color='r',linestyle=':')
ax.legend()
ax.set_title('Histogram of Total Candies per Handful')
```

```python
fig, ax = plt.subplots(1)
sns.boxplot(candy_df['Total'], showmeans=True)
ax.set_xlabel('Total')
ax.set_ylabel('Candies per handful')
ax.set_title('Boxplot of Total Candies per Handful')
```

## Intermission: Short Professor-led discussion about questions (5-10 mins)
After everyone has collected data and looked at plots, Professor will lead discussion about the following quesions **(not required to write anything)**.

1. Does anything jump out at you immediately from the data?


**Questions:** Instead of having ~30 different students in class each pull out 1 handful, consider if instead just 1 person had pulled out ~30 handfuls.
1. Why would that not be as good an estimate of an average person's handful?  
2. Would you expect the mean or standard deviation to be any different?  If so, how?

**Questions:** If you know exactly 100 people will visit your house and each take a handful:

3. Using the data you gathered, how much candy would you **expect** to need?
4. If you buy exactly that much, based on your data very roughly what is the chance that you'll run out of candy before everyone gets some?
5. What factors might be different if you were actually giving out candy on Halloween vs. how we collected data?

## Part 3: Comparing Distributions of a Categorical Variable (10 mins)
Now, pretend that the bag of candy is no longer available to you so you can't examine it.  (This kind of situation happens often in real life because it can be difficult to quantify the exact conditions for which an experiment was conducted.)  

1. Do you think the distribution of the types of candy in the bag could potentially influence your handful results?  If so, how? **(1 point)**

*Write your answer here*  

Suppose you decide you would like to know the distribution of candy that was in the bag, but unfortunately you no longer have to bag to count them exactly.  If you assume the likelihood of a candy being in a handful is the same as the likelihood a candy would be randomly selected from the bag, you might **estimate** the bag distribution using your collected data by computing the total proportion of each candy.  Complete the code below for the other 2 candies: **(1 point)**

```python
kisses_proportion = candy_df['Kisses'].sum()/candy_df['Total'].sum()
starbursts_proportion = ...
miniatures_proportion = ...
```

I actually already counted all the candies in the bag, and the true distribution is 50, 50, 50.

**Questions (1 point):**  
1. What kind of distribution is this?
2. What is the actual proportion of each candy in the bag?

*Write your answers here*

Now, calclulate the **Total Variation Distance** for your measured candy proportions to the true distribution. **(1 point)**

To compute the *Total Variation Distance (TVD)* between two distributions:
* For each category, subtract the two proportions
* Compute the absolute value of each differences
* Add up those values
* Divide by 2

For example, TVD between two distributions [0.5, 0.5] and [0.3, 0.7] is
* (abs(0.5 - 0.3) + abs(0.5 - 0.7))/2 = (0.2 + 0.2)/2 = 0.2

```python
TVD = ...
```

If TVD is close to 0, it means the distributions are very similar.  If TVD is close to 1, it means distributions are very dissimilar. How close was our estimate?

## Part 4: Hypothesis testing (15 mins)
You decide to investigate if **miniatures are actually selected at a different rate than other candy when grabbing a handful**.  And you realize, now that you know the true distribution of candy types, you can use your collected data to evaluate this by using a hypothesis testing process!

Write a null hypothesis and an alternate hypothesis that you could use for determining if there is evidence that miniatures are selected at a **higher** rate than would be predicted by random chance. **(2 points)**


* **Null Hypothesis:** ...
* **Alternate Hypothesis:** ...
* **Statistic you'll use to evaluate:** ...

Now, complete the code blocks below to simulate your statistic "under the null". **(3 points)**

```python
# first, set t1 to the statistic you will use
observed_t1 = ...
```

```python
# 1. Make a lot of simulated data under the null hypothesis
# 2. Calculate the statistics t1 from simulated data

statistic = pd.DataFrame(columns=['t1'])

for simulation in range(1000):
    num_candies_actually_sampled = ...
    simulated_candy = pd.DataFrame(random.choices(['M','K','S'], k=num_candies_actually_sampled), columns=['candy'])
    # minatures should be a True, False array indicating which candies are miniatures
    miniatures = ...
    # compute the statistic from the simulated data
    sim_t1 = ...
    statistic.loc[len(statistic)] = [sim_t1]
statistic
```

```python
# 3. Plot empirical distribution of t1 (under the null)

fig, ax = plt.subplots(1)
sns.histplot(data=statistic, x='t1', stat='density', label='Under the null')
# write code to plot the observed statistic t1
...
ax.legend()
```

From this plot, how much evidence do you see "against the null"?  Which region of the graph is the p-value region?  Write the code to compute the p-value below. **(1 point)**

```python
# first, set t1 to the statistic you will use
p_value = ...
```

## BONUS
If you would have instead decided you wanted to see if miniatures are selected at a **different** rate than would be predicted by random chance, what would you instead use for the statistic?

*write your answers here*

For extra practice, repeat Part 4 of the lab using this new statistic instead.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 13

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert my_guess > 0
```
