```python
# Initialize Otter
import otter
grader = otter.Notebook("lab1.ipynb")
```

# Lab1: Fun with Markdown

This is a Jupyter Notebook, which is an interactive “notebook” where we will be writing a mixture of Python code and text explanations of our results. Each of these boxes is a cell, which can contain a piece of code or text. You’ll become very familiar with it over the semester.


# 1. Hello Team

Same as last time, find your assigned lab partner, sit next to them, and introduce yourself (but no vegetables this time).  Remember, even though you both need to do the lab on your own laptop and submit it individually, please collaborate! 

1. Your name. 
2. Your year and area or planned area of study at UNC. 
3. Why you are interested in Data Science.

That's the last time labdswill explicitly call out this step, but please do it every time!  From now on, section one will just be:

Edit the code cell below with your name and onyen and partner name and onyen.  Then execute the code cell and the autograder test.  Be sure to run the code cell at the top of the file to initialize the autograder!

```python
# edit these names
me = ["Rick Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

```python
grader.check("q0")
```

# 2. Hello Comments
What is a Comment in Code?
A comment is text in your code that Python ignores when running the program.
Comments are for humans to read, not for the computer.
You can use them to explain what your code does, leave notes for yourself, or make the code easier to understand for others.

In Python, a comment starts with a # symbol.
Anything after # on the same line will not be executed.
Example:

```
# This line is a comment. It will not run.
```
Often comments are added to the end of lines of code to explain that line of code.  For example:
```
print("Hello, world!")  # This prints a greeting
```
Change the comment in the code cell below to correctly match what the code does (execute the code cell to see what it does first!)

```python
print(me[1]) # This prints my name. 
```

# 3. Hello Markdown

Throughout DATA110, you will frequently be asked to write explanatory text in which you describe your problem-solving strategy and the structure of your code. One way to do this will be in code comments. Another important and more expressive way is with *Markdown cells.* Markdown (https://en.wikipedia.org/wiki/Markdown#Example) is a simple, readable way to create formatted text. It was invented by John Gruber and Aaron Swartz. This cheatsheet (https://commonmark.org/help/) contains a useful and concise set of instructions for Markdown editing. 

## 3.1 Simple Markdown
To create a Markdown cell, first create a code cell like you did before in Lab 0. Then, with your cursor in the code cell, click the dropdown menu in the toolbar that says "Code".  Select "Markdown" from this menu instead. 

Return to the cell and type the following characters, exactly as they appear:

`*To boldly go* **where no one has gone before.**`

Then, "run" the Markdown cell (`shift + enter` again) to see your text rendered with italic and bold formatting. 

## 3.2 More involved Markdown
Now, try to reproduce in Markdown the following algorithm for enjoying a relaxing cup of tea. 

<br>

<figure class="image" style="width:40%">
  <img src="https://raw.githubusercontent.com/PhilChodrow/PIC16A/master/_images/making-tea.png" alt="A list of six items describing how to make tea. The list reads:             
    1. Boil water.
    2. Place tea bag in mug.
    3. Pour water in mug. **Be careful, water is hot!**
    3. *Optional: add milk and/or sugar*. 
    4. Drink tea. 
        - `while cup is not empty`
            - Carefully raise mug to mouth. 
            - Take a sip. 
    5. Clean mug. ">
  <figcaption><i></i></figcaption>
</figure>


\[add Markdown cell immediately below\]

## 3.3 Let's get visual
Sometimes, you might find it useful to include images in Markdown cells.  Create a markdown cell below this cell and paste the following text:

```![AltText](https://media.giphy.com/media/vFKqnCdLPNOKc/giphy.gif)```

Try adding your own images.

## 3.4  Let's get connected
It is also easy to add general links into markdown.  Try creating a Markdown cell below this one and add the following text to it:

```[Markdown basic syntax](https://www.markdownguide.org/basic-syntax/)```

There are many other things you can do in Markdown.  For a good summary, see:

# 5. Goodbye Lab 1

Great job! It's time to submit the worksheet. Make sure to run all cells (you can click on the fastforward button above (>>)), download the zip file, and submit it on Gradescope under `Lab 1`.  


```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Gradescope Lab1

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q0 (points: 3)
```
>>> assert me[0] != 'Rick Marks'
```
