```python
# Initialize Otter
import otter
grader = otter.Notebook("lab2.ipynb")
```

# Lab 2: Variables and data types


```python
# write your code here
me = ["Rick Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

```python
grader.check("q0")
```

# 1. Data collection (10 minutes)

We'll start today by collecting survey data from 4 classmates + yourself. We'll represent our collected data as a table, which you can write in Markdown as below. Here's a quick guide to Markdown tables (https://www.codecademy.com/resources/docs/markdown/tables), and you can also double click on the cell below to see how I wrote it.

| # | First name | What vegetable would you be? | Shoe size | Birthday month | (Question of your choice) |
|---| -----------| -----------| -----------| -----------| -----------| 
|1 | | | | | 
|2 | | | | | 
|3 | | | | | 
|4 | | | | | 
|5 | | | | | 


Please mingle (try to talk to someone you haven't talked to yet!) and fill out the table. For example, here's one row of data that I have gathered at home:

| # | First name | What vegetable would you be? | Shoe size | Birthday month | Are you a cat? |
|---| -----------| -----------| -----------| -----------| -----------| 
|1 | Puppycat | Corn | 0 | July | True |

Fill in the blank with either `row` or `column`:

- Each *classmate* corresponds to one _____.
- Each *question* corresponds to one _____.

# 2. Data types (5 minutes)

In last lecture, we learned four different data types in Python: `int`, `float`, `str`, and `bool`. If you wanted to work with these data using Python, which types would be the best fit and why? Would you use the same type for all of them? Discuss with your group.

Fill in your guess for the best data type for each piece of information. 

- First name:
- Vegetable:
- Shoe size:
- Birthday month:
- Cat? (or replace with your question):


Discuss your answers with your partner. There can be multiple valid answers. For example, `'July'` or `7` are both valid ways to talks about birthday months. 

# 3. Name (variable) assignment (5 minutes)

Now that you have made a decision on data types, let's practice using it. You can go back to your small group to work on this together.
1. Pick one row from your table.
2. In the cell below, assign the first name of the individual to the variable `name`. For example, `name = 'Puppycat'`.
3. In the next cell below, check the type of the variable `name`. Do you remember what function to call?
4. Repeat steps 2-3 with the remaining four columns. Add new cells as needed. Try to pay attention to the data type as well as what a descriptive variable name would be.

```python
# Write your code here
name = ...
```

```python
# Write your code for other variables here

```

```python
grader.check("q1")
```

# 4. Bag of Words (15 minutes)

**Note**: Some questions in the labs will often use programming elements beyond what we've learned so far.  Don't worry if you have trouble understanding the code in these questions! We do not expect you to understand it at this point. We just want to expose you to new concepts and help you learn how to guess what a function does by running it on different inputs. You are **not** expected to be an expert in this material or be able to write this code in an exam.

In class we learned about **Bag-of-Words model** for converting text to numbers.  A Bag-of-Words is a representation of text that describes the occurrence of words within a document. It involves two things:  
1. A vocabulary of words.
2. A count of how many times a word occurs.

It is called a “bag” of words because the order of the words is discarded. Bag-of-Words does not represent where the words occur in the document, only their count.

In this exercise, Bag-of-Words is implemented using only basic string operations.  There are higher level operations from some Python packages that could make this easier, and more efficient, but for this exercise, we won't be using them.

Our Bag-of-Words result will be represented as a **Table**.  Each row (data sample) will correspond to a document (a string, in our case) and each column (feature) will correspond to a word in the vocabulary.  Each cell in the table will be the count of the word in the document. 

Run the following cell as is. There's no need to change it or understand it. But if you're curious, the Bonus section at the end will help.

```python
import pandas as pd
import numpy as np

def empty_BoW():
    return pd.DataFrame({'Document':[]})

# process a particular document sample for Bag of Words (add new words to vocabulary, compute word counts)
def update_BoW_for_rownum( table: pd.DataFrame, rownum: int ):
    
    # set all counts for the row to zero
    num_words_in_vocabulary = table.shape[1] - 1
    table.loc[rownum] = [table.at[rownum,'Document']] + [0] * num_words_in_vocabulary
    
    # process each word in the document
    for word in table.at[rownum,'Document'].split():
        
        # check if word is already in the vocabulary
        if (word in table.columns):
            
            # if so, increment count for this word
            table.at[rownum,word] = table.at[rownum,word] + 1
        else:
            # otherwise, add word to vocabulary and set count for all rows to zero
            num_rows = table.shape[0]
            table[word] = [0] * num_rows
            
            # then set count for new word for this row to 1
            table.at[rownum,word] = 1

# add a new document to the Bag of Words
def add_document_to_BoW( table: pd.DataFrame, doc: str ):
    num_words_in_vocabulary = table.shape[1] - 1
    new_document_num = table.shape[0]
    
    # add a new row that is the document plus 0 counts for every word in vocabulary
    table.loc[new_document_num] = [doc] + [0]*num_words_in_vocabulary
    
    # update the counts for the new row
    update_BoW_for_rownum(table,new_document_num)
```

First, lets create an empty  Bag-of-Words:

```python
bow_table = empty_BoW()
bow_table
```

And then start adding documents (sentences) to it:

```python
add_document_to_BoW(bow_table,'Harlin loves cats but dislikes dogs')
bow_table
```

```python
add_document_to_BoW(bow_table,'Rick loves dogs and cats')
bow_table
```

```python
add_document_to_BoW(bow_table,'Harlin dislikes cats but loves dogs')
bow_table
```

What do you notice that is interesting about row 0 and row 2?

[Write what you think here]

Try adding some documents of your own to see what happens.

```python
# Uncomment the next few lines, replace each ... with a new sentence string and run this cell

# add_document_to_BoW(bow_table, ...)
# add_document_to_BoW(bow_table, ...)
# bow_table
```

# Bonus

The code above has several things that are likely unfamiliar, but by experimenting you can probably figure some of them out.  To help understand parts of the code above, run each of the follow cells to learn what they do, and discuss with your group:

```python
[0] * 10
```

```python
s = 'this is strange'
s.split()
```

```python
table = pd.DataFrame({'name':['Harlin','Rick','Puppycat','Daisy'],
                       'col1': [11, 12, 13, 14], 
                       'col2': [15, 16, 17, 18]})
table
```

```python
table.shape # Is this (no. of rows, no. of columns) or (no. of columns, no. of rows)?
```

```python
table['name'] # Is this a row or a column?
```

```python
table.loc[2] # Is this a row or a column?
```

```python
table.at[2,'name']
```

# 5. Goodbye Lab 2

Great job! It's time to submit the notebook. Make sure to run all cells (you can click on the fastforward button above (>>)), download the zip file, and submit it on Gradescope under `Lab 2`.  


```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Gradescope Lab2

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q0 (points: 1)
```
>>> assert me[0] != 'Rick Marks'
```


Tests for q1 (points: 4)
```
>>> assert type(name) == str
```
