```python
# Initialize Otter
import otter
grader = otter.Notebook("lab0.ipynb")
```

# Lab 0: Welcome (10 points)

Welcome to your first data science lab!
In this lab, we will learn how to use **Jupyter Notebooks** and write our very first lines of Python code.  
No programming experience? No problem! We’ll go step by step. We'll spend the next few weeks explaining Python programming. For now, this is just to whet your appetite and get you familiar with the tools we are using in this class.

**Goals for Today:**
1. Meet a classmate.
2. Learn how to run code in Jupyter.
3. Practice writing and running simple Python commands and Markdown text.
This is a Jupyter Notebook, which is an interactive “notebook” where we will be writing a mixture of Python code and text explanations of our results. Each of these boxes is a cell, which can contain a piece of code or text. You’ll become very familiar with it over the semester.

# 1. Hello Team

As mentioned in lecture, find the person you are paired with for today and sit beside them.  You will be paired with a different person for each lab.  Labs are meant to be collaborative, so please communicate with your teammate.

Introduce yourselves:

1. Your name. 
2. Your year and area or planned area of study at UNC. 
3. Why you are interested in Data Science. 

Once each of you have shared your interest in Data Science move on to the following question. You might want to spend a minute or so thinking about this one.  

> *If you were a vegetable, which vegetable would you be, **and why**?* 

Make sure that you have both selected a vegetable and a reason. Note that this question is **not** the same as "what is your favorite vegetable?"

For the purposes of this question, botanical vegetables and culinary vegetables are valid choices. For example, "tomato" or "squash" are both fully acceptable responses.


## 2. Hello Jupyter

This is a Jupyter Notebook, which is an interactive “notebook” where we will be writing a mixture of Python code and text explanations of our results. Each of these boxes is a cell, which can contain a piece of code or text. You’ll become very familiar with it over the semester.

- **Markdown cells** (like this one) contain text, which can be prettily formatted.   
- **Code cells** contain Python code you can run.

Let's start by running code that’s already written for you.

Below is a code cell. To run it:  
- Click the cell.
- Press **Shift + Enter** (Windows/Mac) or use the **Run ▶** button in the toolbar.

You should see some text appear *below* the cell.

```python
print("to boldly go where no one has gone before")
```

### Now You Try

In the code cell below, change the text inside the quotation marks (`" "`), and then run it again.

Example:  
```python
print("I love dogs and pizza!")
```


```python
print("Replace this with your own message")
```

An important feature of the Jupyter interface is that the Python environment is persistent between code cells. That means that variables that you define in one code cell can be used in later code cells. 

In the code cell below, type: 

```python
vegetables = ["sweet potato", "Carolina reaper"]
```

(or replace the vegetables with the vegetables you each chose in 1). 

Run the code cell (`shift + enter`). If you get any kind of error, it is likely a punctuation or formatting problem.

```python
# write your code here
vegetables = ...
```

```python
grader.check("q2")
```

The variable `vegetables` is now part of the Python session and is available for future use. For example, type the following code in the cell below and run it. 

```python
print("If we were a vegetable, we would be", vegetables, ".")
```

You'll learn more about how this works in the next week or so. 

```python
# write your code here

```

Sometimes, you'll need to write some code, but there won't be a code cell present. In this case, you can use the `+` symbol on the toolbar to add a new code cell. Add a new cell now, immediately below this text, and type the following code: 

```python
print("One of us would be a", vegetables[0], ".")
```

\[add code cell immediately below\]

On occasion, you may find it useful to *restart the kernel*. Restarting the kernel has the effect of erasing all variables in the current Python session. 

**Restart the kernel by clicking the ↻ button in the Jupyter toolbar.** If asked whether you are sure you want to restart the kernel, choose "yes" 

Create a code cell below, and type in the following code: 

```python
print("And some of us would be", vegetables[1:], ".")
```

\[add code cell immediately below\]

You should observe a `NameError: name 'vegetables' is not defined`. This is because, after restarting the kernel, the variable `vegetables` is no longer present in our namespace. To get it back, go back to the code cell in which you first defined this variable and run it. Then, the cell above will run without error. 

The toolbar also contains useful tools for duplicating cells, removing cells, and moving cells up and down.  Feel free to experiment.

# 3. Howdy partner
Edit the code cell below with your full name and onyen, as well as your partners:

```python
# write your code here
me = ["Rick Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

Run the cell below to make sure everything looks correct.

```python
print("My name is", me[0], "and my onyen is", me[1])
print("My esteemed partner's name is", partner[0], "and their onyen is", partner[1])

```

# 4. Hello autograder
You might have noticed the code cell above in section 2 that said ```grader.check("q2")```  This is the autograder; it runs specific instructor-created tests to make sure your code is correct.  Some of these tests (non-hidden) will give immediate feedback when you run the cell block; those tests are designed to help you fix your code as you are writing it.  Others tests (called "hidden" tests) do not provide feedback and are intended for evaluation.  When you run an autograder cell, you will only find out if you successfully pass the non-hidden tests.  For DATA110H, you will be able to see the scoring result of hidden tests only when you submit your solution to Gradescope.

The following test has 1 non-hidden and 2 hidden tests.  Go ahead and run it:
  

```python
grader.check("q4")
```

If you've updated ```me``` in the code above, the non-hidden test will pass.  For educational purposes, I've designed the hidden tests so the first will pass and the second will fail.  After following the submission instructions below, take a look at the Gradescope results and make sure you understand them; we'll be following this process for nearly every lab and homework assignment all semester. 

# 5. Goodbye first (and zeroeth) recitation

Great job! It's time to submit your work to Gradescope.
After you've submitted it successfully, feel free to play around with Jupyter some more.



### Step 1 – Run All Cells and Save
- My favorite thing to do in Jupyter Notebook is to click the ⏩️ button on the toolbar. This refreshes the memory, then runs all cells in order.
- Then, click **File → Save Notebook** to make sure your notebook file (`.ipynb`) is up to date. Or run `command (or control) + s`.

### Step 2 – Export to PDF
  1. Choose **File → Print**.
  2. In your browser’s print dialog, change the **Destination** to “Save as PDF.”
  3. Click **Save** and name the file clearly (e.g., `lab0.pdf`).

### Step 3 - Run the Last Cell
1. Run the last cell (below)
2. Download the zip file to your computer.

### Step 4 – Submit to Gradescope
- Log into **Gradescope** and open the Lab 0 assignment. You can get there from Canvas assignment as well.
- Upload **both files** by highlighting both items and dragging and dropping them to the submission box:
  1. Your **`.zip` file** (generated by otter).
  2. Your **`.pdf` file** (exported from the notebook).
- Make sure your submission is **complete and readable** before clicking submit.

### Step 5 – Confirm Your Submission
- After submitting, check Gradescope’s preview to ensure your files display properly.
- If anything is missing or unreadable, re-export and resubmit before the deadline.


```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Lab0

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2 (points: 6)
```
>>> assert type(vegetables) == list
```


Tests for q4 (points: 6)
```
>>> assert me[0] != 'Rick Marks'
```
