```python
# Initialize Otter
import otter
grader = otter.Notebook("lab3.ipynb")
```

# Lab3: Control flow

This activity will be on writing Python code that can test whether a password is valid or not. We'll use this time to practice writing variable assignments, control flow (if, for) and thinking through logic.


**Submission instruction**: Please generate a pdf file (File-> Print or cmd + P in Mac) and a zip fle and submit both to Gradescope.

**This exercise can be challenging, so definitely talk through the logic before trying to write it down in code.**


```python
# write your code here
me = ["Rick Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
name=age=has_valid_length=0
```

```python
grader.check("q0")
```

# 1. Getting user inputs (5 minutes)

We haven't shown you this in class yet, but there's a nice function called `input` in Python that can ask for and save user input. Copy paste and run the code below to see an example.

```Python
name = input("What is your name?: ")
name
```

```python
# paste and run the code here!

```

What happened here? In the first line, we asked the user for their input with the prompt "What is your name?: ", and saved their response to a variable `name`. In the second line, we checked to see what's saved in `name`. Make sure you understand these two lines.

**Q**: What is the data type of the variable `name`? Write Python code that returns the type of the variable `name` and assigns it to the variable `name_type`.


```python
# Write your code here
name_type = ...
```

```python
grader.check("q1")
```

Let's try this one more time.

```Python
age = input("What is your age?: ")
print(f'{name} is {age} years old.')
```

```python
# paste and run the code here!

```

What did you notice? The first line looks very similar to the earlier code, but this time we ask the user for their age and save it in a new variable called `age`. 
The second line `print`s out a full sentence using both the variable `name` and `age`. 

**Q**: What is the data type of vaiable `age`? Write Python code that returns the type of the variable `age` and assigns it to the variable `age_type`. Is that what you expected?

```python
# Write your code here
age_type = ...
```

```python
grader.check("q2")
```

**Remark**: The `f` in the beginning of the string tells Python that this is a `f`ormatted string, and that whenever there are curly brackets `{ }`, they should replace that variable with its value. In other words, the curly brackets are placeholders. If this is confusing, copy paste the code below to a new code cell, and try running it to see the difference with the earlier print statement. If you're curious, take a look at [this tutorial](https://www.w3schools.com/python/python_string_formatting.asp) or google 'f-string python'. You are not expected to memorize this syntax or will ever be tested on this, so don't stress out too much.

```python
print('name is age years old.')
```



**Q**: Now we have the tools to ask the user for their password. Write one line of code below that asks the user for `input()` on their password with the prompt "Enter your password: ", and assigns the response to a new variable named `password`. Don't use your actual passwords!!

```python
# Write your code here!


```

# 2. Testing for length (5 minutes)
**Q**: We'll start by writing code to test whether the password is at least 8 letters long. If it is, we will assign variable `has_valid_length` to `True`. If the password is too short, we will assign variable `has_valid_length` to `False`. Refer back to lecture and talk to your group if you are confused. 

**Q**: What data type the variable `has_valid_length` be? (no need to write code to test this)

**Hint**: What built-in Python function can we use to easily count the number of characters of a string?  In other words, the length of the string?


```python
# Write your code here
...
```

```python
grader.check("q3")
```

# 3. Testing for digits (15 minutes)

Often websites require passwords to have at least one number or digit. There are multiple ways to achieve this through Python, and I'll walk you through one such method. I'm going to count the total number of digits in the password, and if that value is greater than 0, I'll declare that password to have passed this test. 

**If you're confused**: Definitely talk through the logic with your group before trying to write Python code. If it's still confusing - don't stay stuck! TAs are ready to help you. They can also share starter code (or see part 5 below for spoilers).

Here's a step by step process.

1. We are going to start by assigning a new variable `num_digits` to integer 0, which we hope (by the end of our code) will be equal to the total number of digits in the password. For example, if my password is 'puppycat999', num_digits should be 3 at the end our code.
   
2. Then, we will **iterate** through every `letter` in the `password`.

**Hint**: Should we use an `if`-`else` statement or a `for` loop?

3. While looking at every `letter`, let's check if it is a digit. You can do this in two ways. You can check whether `letter` is `in` '0123456789', or use the string method `isdigit()`, as in `letter.isdigit()`. See slides to remember how to use string methods.

 **Hint**: Try running this code (and play around with the letters) to understand this step.

```python
letter = 'l'
print(letter.isdigit())
print(letter in '0123456789')

letter = '5'
print(letter.isdigit())
print(letter in '0123456789')
```

4. If the `letter` is a digit, increase the value of `num_digits` by 1.

5. Once we've investigated every `letter` in the `password`, assign `has_digits` to be `True` if `num_digits` is greater than 0, otherwise `False`.


**Q**: Once you talked through the logic and have an idea of what to do, try writing the code below. My solution has 8 lines.

```python
# Write your code here!


```

# 4. Bonus Combining two criteria

Now that we have checked two criteria and saved those results in two `bool` variables `has_valid_length` and `has_digits`, we can combine those information to tell the user whether their password is valid or not. We will do this with an `if`-`else` statement.

Run the following cell, which does the following:

1. If `has_valid_length` **and** `has_digits`, print 'Password is valid!'. We'll discuss **and** at a later lecture.
2. Otherwise, print 'Password is not valid. Please try something else.'.

```python
# Run this cell!

if has_valid_length and has_digits:
    print('Password is valid')
else:
    print('Password is not valid. Please try something else.')

```

# 5. Bonus Putting everything together [spoilers/hints]

You won't have to write anything new for this part. 

**Q**: Copy paste the bits and pieces you've already written to create a full program. For example, this is what my code structure looks like.

```Python
# ask user for their password
password = ...

# check if password contains at least 8 letters
if ... >= 8:
    has_valid_length = ...
else:
    has_valid_length = ...

# check if password contains at least one number
num_digits = ...

for ... in ... :
    if ... :
        num_digits = num_digits + 1

if ... :
    has_digits =  ...
else:
    has_digits =  ...

# combining two criteria
if has_valid_length and has_digits:
    print('Password is valid')
else:
    print('Password is not valid. Please try something else.')
```
**Q**: Is 'password' a valid password?

```python
# paste and run code here! Try running it multiple times with different password to check that it works.


```

# 6. Bonus Bonus

In reality, password checkers have more stringent requirements. If you feel up to the challenge and want more practice, try adding: 

- Password should contain at least one upper case letter.
- Password should contain at least one lower case letter.
- Password should contain at least one special character ~`!@#$%^&*()-_+={}[]|\;:"<>,./?

**Hint**: If you focus on the big picture logic, all of them should feel very similar to one of the two we have already done. Which one?

**Challenge**: Can you write it with just one for loop?

```python
# Write your code here


```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Gradescope Lab3

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q0 (points: 1)
```
>>> assert me[0] != 'Rick Marks'
```


Tests for q1 (points: 2)
```
>>> assert name_type == str
```


Tests for q2 (points: 2)
```
>>> assert age_type == str
```


Tests for q3 (points: 2)
```
>>> assert type(has_valid_length) == bool
```
