```python
# Initialize Otter
import otter
grader = otter.Notebook("lab6.ipynb")
```

# Lab 6 Defining functions
In this activity, our focus will be on writing some efficient Python functions for checking whether numbers are prime. 


```python
# write your code here
me = ["Rick Marks", "rlmarks"]
partner = ["Piper Marks", "piper"]
```

```python
grader.check("yourname")
```

# 1. Test even and odd numbers

Run the cells below. What do they tell us about the number 9?

**Hint**: Recall that `n % i` is the remainder of dividing `n` by `i`.

```python
# Run this code
# This is an expression which is evaluated to a data object
9 % 2 == 0 
```

```python
# Run this code
# What does == check for?
9 % 3 == 0 
```

```python
# What is the type of this object?
type(9 % 2 == 0)
```

[Write your observation here]

Now let's turn that into a function so that we can test multiple numbers, not just 9. Write a function called `is_even` that takes in one integer parameter `n` and returns `True` or `False` depending on whether `n` is an even number or not.

```python
def is_even(n):
    ...
```

```python
grader.check("is_even")
```

Write a new function called `evens_up_to` that takes in parameter `n` as before, but now returns a `list` of all even numbers between 1 and `n` (inclusive).

For example, `evens_up_to(5) = [2, 4]` and `evens_up_to(8) = [2, 4, 6, 8]`.

**Hint**: Consider using a `for` loop with `range`. 

**Hint**: You'll find a very helpful example in lecture 5 slides.

```python
def evens_up_to(n):
    ...
```

```python
grader.check("evens_up_to")
```

# 2. Test Primes

We are now going to try something more advanced.

Write a function called `is_prime(n)` that tests whether an integer `n` is prime. Do so by checking whether `i` divides `n` for all integers `i` where $2 \le i < n$. You should return `True` if the input is prime, and `False` otherwise. 

**Hint**: Remember that `0` and `1` are not prime. You can safely assume that the input $n \ge 2$. 

**Hint**: Recall from lecture 5 that once the function `return`s something, no more code inside the function gets run.

```python
# write is_prime(n) here

...
```

```python
grader.check("is_prime")
```

**Note on grading**: If you have gotten to this point, you will get full credit for today's lab! It is not necessary to complete the entire worksheet in order to receive full credit, provided that you are actively working on the activity for the entirety of the Lab period. If, however, you would like to meet up later with your partners to finish up the activity, then please do -- this will be very helpful for your learning (and therefore your performance on exams). 

# 3. Test primes faster

You may remember from math class that it's not necessary to check *all* numbers `i < n` to tell whether `n` is prime. In fact, it suffices to check only numbers `i` such that $i \leq \sqrt{n}$ (why?). Write a new function called `fast_is_prime(n)` that takes advantage of this fact. 

Check your function using the same tests as above. 

***Hint***: You can get square root of 2 by running the following code:

```python
import math
k = math.sqrt(2)
print(k)
```

**Hint**: You should be able to accomplish this by changing just one line from your previous function.

```python
import math
```

```python
# write fast_is_prime(n) here

...

```

```python
grader.check("fast_is_prime")
```

# BONUS BONUS. Compare Speed

Jupyter comes with a special "magic macro" that lets you compare the execution speed of your functions. To use it, add `%timeit` ("time it") in front of your function call. Write and run the following two lines: 

```python
%timeit is_prime(9973)
%timeit fast_is_prime(9973)
```

You'll see some summary statistics based on a large number of repetitions of your function. How much faster is `fast_is_prime` than `is_prime`? 

```python
# try testing 9973

```

Now test `n = 1000003`. How much faster is your implementation of `fast_is_prime()`? 

**Note:** What is $\sqrt{1,000,000}$?

```python
# try testing 1000003

```

How long does it take to return all primes up to 9973?

```python
# Run your code here

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 6

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for evens_up_to (points: 3)
```
>>> assert evens_up_to(5) == [2, 4]
>>> assert evens_up_to(10) == [2, 4, 6, 8, 10]
```


Tests for fast_is_prime (points: 1)
```
>>> assert fast_is_prime(3)
>>> assert not fast_is_prime(4)
>>> assert fast_is_prime(9973)
>>> assert not fast_is_prime(9977)
```


Tests for is_even (points: 2)
```
>>> assert not is_even(4909)
>>> assert is_even(4886)
```


Tests for is_prime (points: 3)
```
>>> assert is_prime(3)
>>> assert not is_prime(4)
>>> assert is_prime(9973)
>>> assert not is_prime(9977)
```


Tests for yourname (points: 1)
```
>>> assert me[0] != 'Rick Marks'
```
