```python
# Initialize Otter
import otter
grader = otter.Notebook("hw5.ipynb")
```

## <u>Name:</u> [Your name] and (your onyen)

I encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 5: Summary Statistics, Box Plots, Causation, Correlation  (50 points)

Please complete this notebook by filling in the cells provided.

!! Please submit the generated **zip** file and a **PDF** of your notebook. You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop. All problems will be manually graded, so no need to submit ipynb or zip file. **But you need to make sure all your code is visible and not cut off in the pdf. If we can't see your answer code for a coding problem, points may be deducted.**

!! Points will also be **deducted** if you do not include your name and Onyen in the Markdown cell above!!


## Problem 1 (15 points) Really?  Burritos again?
Well, kind of.  More about the restaurant ratings this time.

```python
# Run this cell!!

import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
ratings = pd.read_csv("ratings.csv")
burritos_types = pd.read_csv("burritos_types.csv")

```

**Problem 1.1 (7 points)**. Restaurant ratings.

Now we have a different type of plot for restaurant ratings.

![ratings.png](./ratings.png)

1. (2 points) Try to recreate this plot including the axis labels and title. I used the same fontsizes as the earlier questions. No worries if you get different colors or small cosmetic differences. **Hint**: First decide what type of plot this is, and which seaborn function you need to work with. Then decide which table you need to work with. My solution code is just 5 lines. If you're writing something more complex than this or spending more than 5 to 10 minutes on this, you might be misunderstanding the prompt.

```python
# Write your code here!
...

```

2. (2 points) Without doing any calculations or coding, do you notice any patterns in the data, i.e. any relationships between ratings from different platforms? Which platform do you "trust" more (many valid answers) and does this plot change your opinion on whose ratings to trust more (it's okay if it doesn't)? Make sure to justify your answers. There are many valid answers that can get full credit.

_Type your answer here, replacing this text._

3. (2 point) Do you notice anything odd about the Yelp plot? Is it *missing* anything, compared to the other two? Tell us what's missing in this Yelp plot, and why that would be the case. **Hint**: Try running the `describe` function on this table and see if you notice anything interesting that could lead to this observation.


```python
# Write your code here! (if you need it)
...
```

[Write your answer here!]

4. (1 point) Calculate the 3 statistical dispersion measures (discussed in class) for Google. Make sure to show your work with code and/or writing.

```python
# Write your code here! (if you need it)
...
```

[Write your answer here!]

**Problem 1.2 (6 points)**. Burrito costs (v2).

You liked your previous plot so much that you decided to make it again for burritos.

![burritos_cost.png](burritos_cost.png)

1. (1 point) Make this plot, including the axis labels and title. It's okay if there are small cosmetic differences. Fontsizes are the same as earlier. Again, my solution code is only 5 lines, so if yours is much longer or complicated than that, you might be doing something wrong.

```python
# Write your code here. Add as many new cells (code or markdown) as you need.
...
```

2. (2 point) Without doing any calculations or coding, what would you say a typical burrito would cost in this town? Is your answer different from Problem 2.2? **Is the distribution plotted here the same as the distribution plotted in Problem 2.2**? Make sure to justify your answers.


_Type your answer here, replacing this text._

3. (1 point) Because of the outlier way up high, the plot is being squished. Can you find out which menu item from which restaurant is causing this problem? **Hint**: My solution uses 1~2 lines of code.


```python
# Write your code here!
...
```

4. (2 point) What is the average burrito cost **without** this outlier? How much less (rounded to the nearest cent) is this value compared to the average burrito cost **with** the outlier?

```python
# Write your code here!
...
```

**Problem 1.3 (2 points)**. Translating insights into action.

Surprise! You're actually a consultant hired by someone trying to break into the burrito industry in California. You are analyzing this dataset to gain some insight so that you can give good strategy advice to your client. What advice would you give to your client in terms of 1) how much they should charge for each burrito and 2) which reviewing platform (Yelp or Google) they should advertise more heavily on? Why? Make sure to justify your answer by referring to plots and calculations you've made in previous problems. There are many valid answers that would get full credit as long as they are well-justified.

[Write your answer here!]

_Type your answer here, replacing this text._

## Problem 2 (17 points) More ROUSes?!?  They don't even exist!

```python
import pandas as pd
import seaborn as sns
rouses_model_df = pd.read_csv("ROUSes.csv")
```

**Problem 2.1 (3 points).** Now add a new feature (column) to `rouses_model_df` for each of the 3 models below, and fill in the feature with data generated by the model.  Name the features `'Length Model 1'`, `'Length Model 2'`, and `'Length Model 3'`.  Note that the first model has a square root; for this you can use `**0.5`.  Hint: if you are unsure of mathematical operation ordering, parentheses are highly recommended!

  Length Model 1 = 1.08 * $\sqrt{Age}$

  Length Model 2 = 1.0 + 0.128 * $Age$

  Length Model 3 = 1.0 + 0.00229 * $Age^2$



```python
# write your code here, replacing ...
...
```

**Problem 2.2 (2 points).** There are many ways to compare data generated from a model to the actual data.  One of the easiest is to overlay plots of the actual and generated data.  One very simple way to do this is to put multiple scatterplots in the same cell, one after the other; seaborn is smart enough to automatically change the color for each of the plots.  Do this in the below; your plot should look like this: \
![](LengthModel.png)

```python
# write your code here, replacing ...
...
```

**Problem 2.3 (1 points).**  Just by looking at the plots, which of the models would you say is the best fit to the actual data?  

_Type your answer here, replacing this text._

**Problem 2.4 (2 points).** Summary statistics can also be used to compare data.  For the `Length` data and the data generated by each model, use pandas to help you compute the range to the nearst tenth.  Assign the 4 range values to the corresponding variables in the cells below.

```python
length_range = ...
length_range_model1 = ...
length_range_model2 = ...
length_range_model3 = ...
```

```python
grader.check("q2.4")
```

**Problem 2.5 (2 points).** Now do the same for median.  In other words, for the `Length` data and the data generated by each model, use pandas to help you compute the median to the nearst tenth.  Assign the 4 median values to the corresponding variables in the cells below.

```python
length_median = ...
length_median_model1 = ...
length_median_model2 = ...
length_median_model3 = ...
print(length_median, length_median_model1, length_median_model2, length_median_model3)

```

```python
grader.check("q2.5")
```

**Problem 2.6 (2 points).** Considering the values you computed, which summary statistic (range or median) would be more useful in deciding the model that best approximates the data? Why? 

_Type your answer here, replacing this text._

**Problem 2.7 (2 points).** Another very useful visualization tool for comparing data is boxplot.  First, execute the code below:

```python
sns.boxplot(data=rouses_model_df)
```

Unfortunately, you should be able to see the Weight column and the Age column are "dominating" the plot, and we aren't even interested in them right now.  In the cell below, complete the code you should assign to `new_argument` to create boxplots for only the 4 columns related to *Length*.

```python
new_argument = ...
sns.boxplot(data=new_argument)
```

**Problem 2.8 (3 points).** Using any technique you wish, decide which of the models below is the best model for `Weight`.  Show some kind of work or visualization in the Python cell, and write what model matches best in the cell below that.

  Weight Model 1 = 51 * $\sqrt{Age}$

  Weight Model 2 = 6.9 * $Age$

  Weight Model 3 = 0.125 * $Age^2$



```python
# show your work here
...
```

_Type your answer here, replacing this text._

## Problem 3. Correlation (4 points)
An important part of data science is understanding the relationships between data.  **Correlation** is the degree to which data change with one another.  **Causality** implies a direct relationship where changes in one thing directly results in changes to another thing.  Data can be correlated without any causality relationship; in fact, this occurs often when the data both depend on a third variable.  For example, number of shark attacks and ice cream sales are often correlated, but one does not cause the other; they both depend on weather conditions.  

**Problem 3.1 (1 point).** Run the code below and observe the plot.  Does the ***y*** value appear to be correlated with the ***x*** value?


```python
import seaborn as sns
x_data = [0,1,2,3,4,5,6,7,8,9,10]
y_data = [1,4,7,10,2,5,8,9,6,3,0]
sns.scatterplot(x=x_data, y=y_data)
```

_Type your answer here, replacing this text._

**Problem 3.2 (1 point).** What would be your estimate for the correlation coefficient ***r***?

_Type your answer here, replacing this text._

**Problem 3.3 (1 point).** Run the code below and observe the plot.  Does the ***y*** value appear to be correlated with the ***x*** value?


```python
x_data = [0,1,2,3,4,5,6,7,8,9,10]
y_data = [-10,-13,-14,-16,-18,-20,-21,-24,-26,-27,-30]
sns.scatterplot(x=x_data, y=y_data)
```

_Type your answer here, replacing this text._

**Problem 3.4 (1 point).** What would be your estimate for the correlation coefficient ***r***?


_Type your answer here, replacing this text._

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework5

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2.4 (points: 2)
```
>>> assert length_range == 6.9
>>> assert length_range_model3 == 6.9
```


Tests for q2.5 (points: 2)
```
>>> assert length_median == 3.9
>>> assert length_median_model3 == 1.2
```
