```python
# Initialize Otter
import otter
grader = otter.Notebook("hw8.ipynb")
```

## Name: [Your name] and (your onyen)

I encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 8: NN and Linear Regression (38 points)

Please complete this notebook by filling in the cells provided.

!! Please submit a zip and a PDF of your notebook. You can make a PDF from JupyterHub by 1. File->Print or 2. Command+P (on Mac), then choose 'Save to PDF'. This saves the file on your laptop. Most problems will be manually graded, but there a couple autograded questions.  No private autograded, so if you pass in JupyterHub, you'll get full Autograder credit on Gradescope. Please make sure all your code is visible and not cut off in the pdf.

!! Points will also be deducted if you do not include your name and Onyen in the Markdown cell above!!

## Problem 1. Are you my neighbor? (13 points)
In this problem, you will be manually implementing the nearest neighbor algorithm in Python using the penguins dataset.  Run the following cell to load the data:

```python
import pandas as pd
penguins = pd.read_csv("palmer_penguins.csv") # Load the penguins dataset from the demo folder
penguins.head()
```

### **Problem 1.1 (2 points)**
Suppose you stumble upon a penguin in a zoo, and you are able to measure it's Culmen Length (49.3) and Culmen Depth (18.2) but nothing else.  In this problem, you will use the nearest neighbor (NN) algorithm to predict that penguin's Sex, Species, and what Island it comes from.  Since you know that nearest neighbor uses distance to find the neighbor, you start by writing a distance function.  Complete the code for the following distance function, replacing ...

```python
# we will need sqrt function, so we import math
import math 

# Compute the distance between two lists of numbers.
def compute_distance(list1, list2, distance_type):
    # list1 (list of float): The first list of numbers.
    # list2 (list of float): The second list of numbers.
    # distance_type (str): The type of distance to compute. Currently only 'euclidean' is supported.
    if (len(list1) != len(list2)):
        raise ValueError("Lists must be of the same length")
    if distance_type == 'euclidean':
        sum = 0
        for i in range(len(list1)):
            ...
        return math.sqrt(sum)
    else:
        raise ValueError("Unsupported distance type: {}".format(distance_type))
    

```

```python
# Test the distance function
print(compute_distance([1, 2], [4, 6], 'euclidean'))  # Should print 5.0

```

```python
# Test the distance function
print(compute_distance([0, 0, 0], [1, 2, 2], 'euclidean'))  # Should print 3.0
```

### **Problem 1.2 (2 points)**
When computing distances, it would be bad if one of the features had missing data.  Similarly, if there is missing data for Sex, Species, or Island we won't be able to give a prediction. Write code to drop all the rows that have NaN for either the features or the things we want to predict, and store the result in a new table called `penguins_clean`

```python
penguins_clean = ...
print(f"Rows before cleaning: {len(penguins)}, Rows after cleaning: {len(penguins_clean)}")
penguins_clean.reset_index(drop=True, inplace=True) # reset index after dropping rows.  This makes sure no index values are skipped.
penguins_clean.head()
```

```python
grader.check("q1.1")
```

### **Problem 1.3 (2 points)**
Finish the code below to find the nearest neighbor of our penguin, replacing the two ... portions.  Use the distance function you finished above. 

```python
def find_zoo_penguin_nearest_neighbor(zoo_penguin_features_list, distance_type):
    nearest_index = -1
    nearest_distance = float('inf') # initialize to a huge distance
    for index in range(len(penguins_clean)):
        penguin_culmen_length = penguins_clean.loc[index, 'Culmen Length (mm)']
        penguin_culmen_depth = penguins_clean.loc[index, 'Culmen Depth (mm)']
        penguin_features_list = [penguin_culmen_length, penguin_culmen_depth]
        distance = ...
        if distance < nearest_distance:
            # update nearest neighbor
            nearest_distance = ...
            nearest_index = index
    print("Nearest neighbor distance:", nearest_distance)
    print("Nearest neighbor species:", penguins.loc[nearest_index, 'Species'].split()[0]) # print only the genus
    print("Nearest neighbor sex:", penguins.loc[nearest_index, 'Sex'])
    print("Nearest neighbor island:", penguins.loc[nearest_index, 'Island'])
    return nearest_index
```

```python
# test the function.  You should get Adelie, FEMALE, Dream as the answer
zoo_penguin = [49.3, 18.2]  # Culmen Length (mm), Culmen Depth (mm)
find_zoo_penguin_nearest_neighbor(zoo_penguin, 'euclidean')
```

### **Problem 1.4 (1 points)**
Oops!  You misread your handwriting.  The Culmen Length was actually 29.3, not 49.3.  Try again.  Good thing you implemented nearest neighbor as a function!

```python
zoo_penguin = ...
find_zoo_penguin_nearest_neighbor(zoo_penguin, 'euclidean')
```

```python
grader.check("q2")
```

### **Problem 1.5 (2 points)**
Copy your distance function from above, but re-write it to also handle the "manhattan" distance.

```python
...
```

```python
zoo_penguin = [29.3, 18.2]  # Culmen Length (mm), Culmen Depth (mm)
find_zoo_penguin_nearest_neighbor(zoo_penguin, 'manhattan')
```

### **Problem 1.6 (1 points)**
Are there any differences in the nearest neighbor results when using Euclidean distance versus Manhattan distance?

*Write your answer here*

### **Problem 1.7 (2 points)**
After many embarrassing attempts, you finally get a measurement of the zoo penguin's flipper length: 217.0.  In the cell below, finish a new function named `new_find_zoo_penguin_nearest_neighbor` that uses Flipper Length as a feature along with Culmen Length and Culmen Depth.  The code will be very similar to the original, so you might want to copy a lot of the code from above.  Make sure your new function prints the prediction results like the old one did.  Don't worry about re-cleaning the data; all the flipper length measurements are there in penguin_clean. 

```python
def new_find_zoo_penguin_nearest_neighbor(zoo_penguin_features_list, distance_type):
    nearest_index = -1
...
    return nearest_index
```

### **Problem 1.8 (1 points)**
What predictions do you get for the zoo penguin now?  Call your new function below with each of the distance types.

```python
...
```

## Problem 2. Parkinsons progression detection (25 points)

We will reproduce some of the qualitative results from 
the following paper, which I've downloaded for you in the same folder:

    A. Tsanas, M. A. Little, P. E. McSharry and L. O. Ramig, "Accurate Telemonitoring of Parkinson's Disease Progression by Noninvasive Speech Tests," in IEEE Transactions on Biomedical Engineering, vol. 57, no. 4, pp. 884-893, April 2010, doi: 10.1109/TBME.2009.2036000.

The Parkinsons Telemonitoring dataset was downloaded from https://archive.ics.uci.edu/dataset/189/parkinsons+telemonitoring.

**It is not necessary for you to read or understand the paper in order to complete this problem. You only need to look at the Tables and Figures in the paper that I point to.**

For context, this dataset is composed of a range of biomedical voice measurements from 42 people with early-stage Parkinson's disease recruited to a six-month trial of a telemonitoring device for remote symptom progression monitoring. The recordings were automatically captured in the patient's homes.
    
Physical test observations are mapped to a metric specifically designed to follow disease progression, typically the unified Parkinson’s disease rating scale (UPDRS) that reflects the presence and severity of symptoms (but does not quantify their underlying causes). For untreated patients, the **total UPDRS spans the range 0–176**, with 0 representing healthy state and 176 representing total disabilities, and consists of three sections: 1) mentation, behavior, and mood; 2) activities of daily living; and 3) motor. The **motor UPDRS ranges from 0 to 108**, with 0 denoting symptom free and 108 denoting severe motor impairment, and encompasses tasks such as speech, facial expression, tremor, and rigidity. 

Columns in the table contain subject number, subject age, subject gender, time interval from baseline recruitment date, motor UPDRS, total UPDRS, and 16 biomedical voice measures. Each row corresponds to one of 5,875 voice recording from these individuals. **The main aim is to predict the motor and total UPDRS scores from the 16 voice measures.**


```python
parkinsons = pd.read_csv('parkinsons_updrs.csv')
parkinsons
```

### **Problem 2.1 (2 points) Concepts of machine learning**

Given the data format and the description above,
- Is this more of a **regression** problem or a **classification** problem? Pick one closest answer.
- What do the **samples** correspond to, and how many are there?
- What are the **target variables (labels, y)**, and how many are there?
- What are the **predictor variables (features, X)**, and how many are there?

*Write your answer here!*

### **Problem 2.2 (6 points) Exploring correlations**

Open the paper pdf file and scroll down to page 887, Table 1. It looks scary, but all it's doing is calculating the **correlations** between each of the voice features with the UPDRS scores. We'll recreate this table step by step.

#### Step 1 (1 point)

Drop the columns `subject#`, `age`, `sex`, `test_time` and save the resulting table to a new variable `parkinsons_small`. Your new table should be of size (5875, 18).

```python
# Step 1
# write your code here
...
print(parkinsons_small.shape)
parkinsons_small.head()
```

#### Step 2 (0 point, it's just copy pasting)

Copy and run the following code, which calculates the correlation between each pairs of columns in `parkinsons_small`. Your new table should be of size (18, 18).

```Python
corrs = parkinsons_small.corr(method='spearman')
print(corrs.shape)
corrs
```

How to read the table: The correlation between `Jitter(%)` and `total_UPDRS` is 0.129237.

Out of scope for DATA110 but just in case you're curious: There are a few mathematical definitions of how to calculate correlation, and Pearson's and Spearman's are two of the most popular methods. The main difference is that Spearman's cares about the ranking of the values, and not the raw values. So it has fewer assumptions on the data than Pearson's and can handle outliers better. But the way the correlation values are interpreted are the same.

```python
# Step 2
# copy code here and run it
...

```

#### Step 3 (2 points)

Since we're only interested in the correlations with the UPDRS scores, 

1. Keep only the columns `motor_UPDRS`,`total_UPDRS` and save the new table to `corrs_table1`.
2. Within the table `corrs_table1`, drop the rows with index `motor_UPDRS`,`total_UPDRS` and save the new table to `corrs_table1` (overwrite it).

```python
# Step 3
# write your code here
...
print(corrs_table1)
```

#### Step 4 (3 points) Interpretation

Compare what you have with the numbers in Table 1, and check if they are **roughly** the same. 

**Note**: The numbers won't be exactly the same, because 5923 recordings are analyzed in the paper, but we only have access to 5875 of them. 

Question: What are your interpretations of these values? Do any of the biomedical voice variables seem highly correlated with the Parkinson's scores? Which variable do you think would be the most important in predicting UPDRS? 

*Write your answer here!*

### **Problem 2.3 (6 points) Prepare data for ML**

Now comes the actual prediction! We will train and test a **linear regression** model to predict motor_UPDRS from the 16 voice variables. In other words, we are assuming that the motor_UPDRS scores are **approximately** equal to

c1\*Jitter(%) + c2\*Jitter(Abs) + c3\*Jitter:RAP + ... + c16\*PPE

for some set of coefficients (c1, c2, c3, ..., c16). And when we train or fit a linear regression model, we are looking for the best set of coefficients that describes the training dataset. When we test the model, we compare the model's predictions with the actual motor_UPDRS scores in the test dataset. If any of this sounds foreign to you, now would be a good time to review the past few lectures.

#### Problem 2.3.1 (2 points)

In an ideal world, we would train our model on the entire dataset, collect data from new patients, then test it on the new data. However, this is obviously not feasible in this case. In cases like this, most people randomly split the *existing* samples into train and test, and "pretend" like the samples in the test set are actually coming from future patients.

We can split the data by patients or by recordings. The authors of this paper chose to split by recordings. 

**Fill in the blank** such that this sentence describes what the code does:

    We will randomly put __1__% of the ___2____ into the __3___ set, and put the remaining __2_____ into the ____3___ set.

- Options for 1: number between 0 and 100
- Options for 2: rows or columns
- Options for 3: train or test

The corresponding code:

```Python
train = parkinsons_small.sample(frac=0.9)
test = parkinsons_small.drop(index=train.index)
```

And then copy and run the code.


*Write your answer here!*

```python
# copy code here and run it
...

```

#### Problem 2.3.2 (4 points)

Given your knowledge of what predictor variables and target variables are, create the following four variables: `X_train`, `y_train`, `X_test`, `y_test`. They are derived from tables `train` and `test`.

Here's one way to check your answer:
```Python
print(X_train.shape, y_train.shape)
    (5288, 16) (5288,)
print(X_test.shape, y_test.shape)
    (587, 16) (587,)
```

```python
# write your code here
...
```

### **Problem 2.4 (2 points) Train and test model**

Copy the code, fill in the blanks (i.e., call with appropriate parameters), and run the code in the next cell.

```Python
from sklearn.linear_model import LinearRegression

lr = LinearRegression(fit_intercept=False)
lr.fit(..., ...)
print(lr.score(..., ...))

coefs = pd.DataFrame(lr.coef_, 
                     index=lr.feature_names_in_, 
                    columns=['Motor UPDRS LR coefficients'])
coefs
```

```python
# Copy the skeleton code here and complete it
...
```

### **Problem 2.5 (4 points) Analyze model results**

#### Problem 2.5.1 (2 points)

Given the $R^2$ score that is printed out, do you think this linear regression model is a good predictive model of motor_UPDRS or not?

*Write your answer here!*

#### Problem 2.5.2 (2 points)
Out of the 16 features, which one(s) are deemed the most *important* or *useful* at predicting motor_UPDRS according to the model? Comment on both the magnitude and direction (negative vs. positive) of the coefficient values and what that means in the context of the problem we're interested in. Are they the same variables you mentioned in Problem 4.2 (based on correlations)? You are welcome to compare your numbers with the reulsts in Table 3 (first column), but not required.

*Write your answers here!*

### **Problem 2.6 (5 points) Repeating the experiment multiple times**

You might say - but Harlin, this is only based on a specific train/test split. What if it just so happened that this set of samples leads to this result? How can we trust our interpretations are generally true? 

You are absolutely right! That's why the authors of this paper repeated this 1000 times to create Table 3. 

Fill in the missing parts of the code below (you shouldn't have to write new code, just copy paste from earlier problems). 

```Python
coefs = pd.DataFrame([], columns=lr.feature_names_in_)

for i in range(1000):
    train = parkinsons_small.sample(frac=0.9)
    test = parkinsons_small.drop(train.index)

    #####
    # Add four lines of code here that creates X_train, y_train, X_test, y_test. 
    #####
    
    lr = LinearRegression(fit_intercept=False)
    # fill in the parameters here
    lr.fit(..., ...)

    coefs.loc[len(coefs)] = lr.coef_

coefs.mean()
```

Then tell us, **how did your analysis on the importance of each feature change?** If it didn't change, that's is also okay.


**Note**: Again, you won't get the same numbers as Table 3 because we don't have access to the full dataset, and there is some randomness involved. But my coefficients except for Jitter (%) looks *qualitatively* similar (I get a 40-ish number for Jitter (%) instead of a -80-ish number, which is odd).

```python
# Copy the skeleton code here and complete it
...

```

```python
# this should give you pretty box plots once you're done with the previous question.
import seaborn as sns
sns.boxplot(coefs, orient='h')
```

*Write your answer here: How did your analysis of the importance of each feature change?*

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Homework8

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.1 (points: 2)
```
>>> assert penguins_clean.shape[0] == 334
```


Tests for q2 (points: 1)
```
>>> assert find_zoo_penguin_nearest_neighbor([29.3, 18.2], 'euclidean') == 136
Nearest neighbor distance: 3.889730067755345
Nearest neighbor species: Adelie
Nearest neighbor sex: FEMALE
Nearest neighbor island: Dream
```
