```python
# Initialize Otter
import otter
grader = otter.Notebook("lab18.ipynb")
```

# Lab 18: k-means (10 pts)
Please work with your final project partner.
  
**Submission instruction**: Please create a zip file and a pdf via File -> Print (or cmd + P on mac), and upload it to Gradescope.   Even if you decide to finish it at home, please submit what you have by the end of class to make sure you get some credit!


```python
# edit these names to your name and your final project partner's name
me = ["Rick Marks"]
partner = ["Piper Marks"]
...
```

```python
grader.check("name")
```

## Part A: Preparing data



Run the following cell to load the penguin dataset as a `pandas` `DataFrame` called `penguins`. I've also supplied code to shorten the penguins species name for convenient exploration and plotting. 

```python
import pandas as pd
import seaborn as sns
from matplotlib import pyplot as plt
pd.set_option("future.no_silent_downcasting", True)

penguins = pd.read_csv("palmer_penguins.csv")

# shorten the species name
penguins["Species"] = penguins["Species"].str.split().str.get(0)
```

For today's first exercise, keep only the following columns: ``'Culmen Length (mm)'`, `'Flipper Length (mm)'`. Calling `penguins.filter([...])` with the column names inside should make this happen. Reassign this table to `penguins`. The updated `penguins` table should have 344 rows and 2 columns.

```python
# Write your code here
...


```

You might have noticed that your table contains rows with `NaN` values. Calling `penguins.dropna()` will remove these rows. Do this below, and reassign the result back to `penguins`. Your updated `penguins` table should have 342 rows and 2 columns.

```python
# Write your code here
...

```

## Part B: Clustering

It's always good to take a look at your data first.  Make a scatterplot of flipper length vs. culmen length.

```python
# write your code here
...

```

How many clusters do you see in the data?

*Your answer here*

Now use k-means clustering to make 3 clusters and label each penguin.  This is just like the demo code from lecture.  I've supplied most if it here:
```Python
from sklearn.cluster import KMeans
kmeans = KMeans(n_clusters=...)
kmeans.fit(...)
print(kmeans.cluster_centers_)
plt.scatter(penguins['Culmen Length (mm)'], penguins['Flipper Length (mm)'],s=20);
plt.scatter(kmeans.cluster_centers_[0], kmeans.cluster_centers_[:,1],  c='red', marker='x', s=100)

```


```python
# copy and complete the code here
...
```

Next, use the trained kmeans to assign each penguin to a cluster.  Again, this is just like the demo code.  I've done most of it for you, just replace the ???
```Python
y_predicted = kmeans.???(penguins)
plt.scatter(penguins['Culmen Length (mm)'], penguins['Flipper Length (mm)'], c=y_predicted, s=20);
```


```python
# copy and finish the code to use the trained kmeans to predict the cluster for each point
...
```

Hmm, maybe the clusters aren't quite what you would expect?  Especially the purple-labeled penguins in the middle left?  Any idea what is going on here?  Discuss with class before moving on.

## Part C: Clustering again

kmeans relies heavily on distances between points, so if there are features with different scales, it can affect the clustering. Use min/max scaling to create a new table called `penguins_normalized`.  This can be done in one line of code!

```python
penguins_normalized = ...
penguins_normalized.head()
```

Repeat the entire clustering process for penguins_normalized.  Be sure to draw the final scatterplot with the labels assigned to each penguin again.

```python
# repeat the entire clustering process for penguins_normalized
...
```

Those clusters should make a lot more sense now!  One final exercise.  Reload the original penguins dataset, and again make a scatterplot of flipper length vs. culmen length, but also use hue to display the species in the scatterplot.

```python
# reload the original penguins dataset, and make a scatterplot of flipper length vs. culmen length and use hue to display the species in the scatterplot.
...
```

How does this compare to your normalized kmeans plot?  What does this imply?

[Write your answer here]

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip and PDF file to Gradescope Lab 18

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for name (points: 1)
```
>>> assert me[0] != 'Rick Marks'
>>> assert partner[0] != 'Piper Marks'
```
